A Novel Rapid Host Cell Entry Pathway Determines Intracellular Fate of Staphylococcus aureus

Author response:

The following is the authors’ response to the original reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The manuscript by Rühling et al analyzes the mode of entry of S. aureus into mammalian cells in culture. The authors propose a novel mechanism of rapid entry that involves the release of calcium from lysosomes via NAADP-stimulated activation of TPC1, which in turn causes lysosomal exocytosis; exocytic release of lysosomal acid sphingomyelinase (ASM) is then envisaged to convert exofacial sphingomyelin to ceramide. These events not only induce the rapid entry of the bacteria into the host cells but are also described to alter the fate of the intracellular S. aureus, facilitating escape from the endocytic vacuole to the cytosol.

Strengths:

The proposed mechanism is novel and could have important biological consequences.

Weaknesses:

Unfortunately, the evidence provided is unconvincing and insufficient to document the multiple, complex steps suggested. In fact, there appear to be numerous internal inconsistencies that detract from the validity of the conclusions, which were reached mostly based on the use of pharmacological agents of imperfect specificity.

We thank the reviewer for the detailed evaluation of our manuscript. We will address the criticism below.

We agree with the reviewer that many of the experiments presented in our study rely on the usage of inhibitors. However, we want to emphasize that the main conclusion (invasion pathway affects the intracellular fate/phagosomal escape) was demonstrated without the use of inhibitors or genetic ablation in two key experiments (Figure5 D/E). These experiments were in line with the results we obtained with inhibitors (amitriptyline [Figure 4D], ARC39, PCK310, [Figure 4C] and Vacuolin-1 [Figure4E]). Importantly, the hypothesis was also supported by another key experiment, in which we showed the intracellular fate of bacteria is affected by removal of SM from the plasma membrane before invasion, but not by removal of SM from phagosomal membranes after bacteria internalization (Figure5A-C). Taken together, we thus believe that the main hypothesis is strongly supported by our data.

Moreover, we either used different inhibitors for the same molecule (ASM was inhibited by ARC39, amitriptyline and PCK310 with similar outcome) or supported our hypothesis with gene-ablated cell pools (TPC1, Syt7, SARM1), as we will point out in more detail below.

Firstly, the release of calcium from lysosomes is not demonstrated. Localized changes in the immediate vicinity of lysosomes need to be measured to ascertain that these organelles are the source of cytosolic calcium changes. In fact, 9-phenantrol, which the authors find to be the most potent inhibitor of invasion and hence of the putative calcium changes, is not a blocker of lysosomal calcium release but instead blocks plasmalemmal TRPM4 channels. On the other hand, invasion is seemingly independent of external calcium. These findings are inconsistent with each other and point to non-specific effects of 9-phenantrol. The fact that ionomycin decreases invasion efficiency is taken as additional evidence of the importance of lysosomal calcium release. It is not clear how these observations support involvement of lysosomal calcium release and exocytosis; in fact treatment with the ionophore should itself have induced lysosomal exocytosis and stimulated, rather than inhibited invasion. Yet, manipulations that increase and others that decrease cytosolic calcium both inhibited invasion.

With respect to lysosomal Ca2+ release, we agree with the reviewer that direct visual demonstration of lysosomal Ca2+ release upon infection will improve the manuscript. We therefore performed live cell imaging to visualize lysosomal Ca2+ release by a previously published method.1 The approach is based on two dextran-coupled fluorophores that were incubated with host cells. The dyes are endocytosed and eventually stain the lysosomes. One of the dyes, Rhod-2, is Ca2+-sensitive and can be used to estimate the lysosomal Ca2+ content. The second dye, AF647, is Ca2+-insensitive and is used to visualize the lysosomes. If the ratio Rhod-2/AF647 within the lysosomes is decreasing, lysosomal Ca2+ release is indicated. We monitored lysosomal Ca2+ content during S. aureus infection with this method (Author response image 1 and Author response video 1). However, the lysosomes are very dynamic, and it is challenging to monitor the fluorescence intensities over time. Thus, quantitative measurements are not possible with our methodology, and we decided to not include these data in the main manuscript. However, one could speculate that lysosomal Ca2+ content in the selected ROI (Author response image 1 and Author response video 1) is decreased upon attachment of S. aureus to the host cells as indicated by a decrease in Rhod-2/AF647 ratio.

Lysosomal Ca2+ imaging during S. aureus infection.

The lysosomes of HuLEC were stained with two dextran-coupled fluorescent dyes. A Ca2+-sensitive dye Rhod-2 as well as Ca2+insensitive AF647. Cells were infected with fluorescent S. aureus JE2 and monitored by live cell imaging (see Author response video 1). The intensity of Rhod-2/AF647 was measured close to a S. aureus-host contact site. Ratio of Rhod-2 vs. AF647 fluorescence intensity was calculated

As to the TRPM4 involvement in S. aureus host cell internalization, it has been reported that TRPM4 is activated by cytosolic Ca2+. However, the channel conducts monovalent cations such as K+ or Na+ but is impermeable for Ca2+ [2, 3]. The following of our observations are supporting this:

i) S. aureus invasion is dependent on intracellular Ca2+, but is independent from extracellular Ca2+ (Figure 1A).

ii) 9-phenantrol treatment reduces S. aureus internalization by host cells, illustrating the dependence of this process on TRPM4 (data removed from the manuscript) . We therefore hypothesize that TRPM4 is activated by Ca2+ released from lysosomes (see above).

TRPM4 is localized to focal adhesions and is connected to actin cytoskeleton[4, 5] – a requisite of host cell entry of S. aureus.[6, 7] This speaks for an important function of TRPM4 in uptake of S. aureus in general, but does not necessarily have to be involved exclusively in the rapid uptake pathway.

TRPM4 itself is not permeable for Ca2+ but is activated by the cation. Thus, it is unlikely to cause lysosomal exocytosis. The stronger bacterial uptake reduction by treatment with 9-phenantrol when compared to Ned19 thus may be caused by the involvement of TRPM4 in additional pathways of S. aureus host cell entry involving that association of TRPM4 with focal adhesions or as pointed out by the reviewer, unspecific side effects of 9-phenantrol that we currently cannot exclude. However, we think that experiments with 9-phenantrol distract from the main story (lysosomal Ca2+ and exocytosis) and might be confusing for the reader. We thus removed all data and discussion concerning 9phenantrol in the revised manuscript.

Regarding the reduced S. aureus invasion after ionomycin treatment, we agree with the reviewer that ionomycin is known to lead to lysosomal exocytosis as was previously shown by others8 as well as our laboratory[9}.

We hypothesized that pretreatment with ionomycin would trigger lysosomal exocytosis and thus would reduce the pool of lysosomes that can undergo exocytosis before host cells are contacted by S. aureus. As a result, we should observe a marked reduction of S. aureus internalization in such “lysosome-depleted cells”, if the lysosomal exocytosis is coupled to bacterial uptake. Our observation of reduced bacterial internalization after ionomycin treatment supports this hypothesis.

However, ionomycin treatment and S. aureus infection of host cells are distinct processes.

While ionomycin results in strong global and non-directional lysosomal exocytosis of all “releasable” lysosomes (~5-10 % of all lysosomes according to previous observations)8, we hypothesize that lysosomal exocytosis upon contact with S. aureus only involves a small proportion of lysosomes at host-bacteria contact sites. This is supported by experiments that demonstrate that ~30% of the lysosomes that are released by ionomycin treatment are exocytosed during S. aureus infection (see below and Figure 2, A-C). We added this new data as well as an according section to the discussion (line 563 ff). Moreover, we moved the data obtained with ionomycin to Figure 2E and described our idea behind this experiment more precisely (line 166 ff).

The proposed role of NAADP is based on the effects of "knocking out" TPC1 and on the pharmacological effects of Ned-19. It is noteworthy that TPC2, rather than TPC1, is generally believed to be the primary TPC isoform of lysosomes. Moreover, the gene ablation accomplished in the TPC1 "knockouts" is only partial and rather unsatisfactory. Definitive conclusions about the role of TPC1 can only be reached with proper, full knockouts. Even the pharmacological approach is unconvincing because the high doses of Ned-19 used should have blocked both TPC isoforms and presumably precluded invasion. Instead, invasion is reduced by only ≈50%. A much greater inhibition was reported using 9-phenantrol, the blocker of plasmalemmal calcium channels. How is the selective involvement of lysosomal TPC1 channels justified?

As to partial gene ablation of TPC1: To avoid clonal variances, we usually perform pool sorting to obtain a cell population that predominantly contains cells -here- deficient in TPC1, but also a small proportion of wildtype cells as seen by the residual TPC1 protein on the Western blot. We observe a significant reduction in bacterial uptake in this cell pool suggesting that the uptake reduction in a pure K.O. population may be even more pronounced.

As to the inhibition by Ned19:

The scale of invasion reduction upon Ned19 treatment (50%, Figure 1B) is comparable with the reduction caused by other compounds that influence the ASM-dependent pathway (such as amitriptyline, ARC39 [Figure 2G], BAPTA-AM [Figure 1A], Vacuolin-1 [Figure 2D], β-toxin [Figure 2L] and ionomycin [Figure 2E]). Further, the partial reduction of invasion is most likely due to the concurrent activity of multiple internalization pathways which are not all targeted by the used compounds and which we briefly discuss in the manuscript.

We agree with the reviewer that Ned19 inhibits TPC1 and TPC2. Since ablation of TPC1 reduced invasion of S. aureus, we concluded that TPC1 is important for S. aureus host cell invasion. We thus agree with the reviewer that a role for TPC2 cannot be excluded. We clarified this in the revised manuscript (Lines 552). It needs to be noted, however, that deficiency in either TPC1 or TPC2 alone was sufficient to prevent Ebola virus infection10, which is in line with our observations.

In order to address the role of TPC2 for this review process, we kindly were gifted TPCN1/TPCN2 double knock-out HeLa cells by Norbert Klugbauer (Freiburg, Germany), which we tested for S. aureus internalization. We found that invasion was reduced in these cell lines supporting a role of lysosomal Ca2+ release in S. aureus host cell entry and a role for both TPC channels (Author response image 2, see end of the document). Since we did not have a single TPCN2 knock-out available we decided to exclude these data from the main manuscript.

Invasion efficiency is reduced in TPC1/TPC2 double K.

O. HeLa cells. Invasion efficiency of S. aureus JE2 was determined in TPC1/TPC2 double K.O. cells after 10 and 30 min. Results were normalized to the parental HeLa WT cell line (set to 100 %).

Invoking an elevation of NAADP as the mediator of calcium release requires measurements of the changes in NAADP concentration in response to the bacteria. This was not performed. Instead, the authors analyzed the possible contribution of putative NAADP-generating systems and reported that the most active of these, CD38, was without effect, while the elimination of SARM1, another potential source of NAADP, had a very modest (≈20%) inhibitory effect that may have been due to clonal variation, which was not ruled out. In view of these data, the conclusion that NAADP is involved in the invasion process seems unwarranted.

Our results from two independent experimental set-ups (Ned19 [Figure 1B] and TPC1 K.O. [Figure 1C & Figure 2N]) indicate the involvement of NAADP in the process. Together with the metabolomics unit at the Biocenter Würzburg, we attempted to measure cellular NAADP levels, however, this proved to be non-trivial and requires further optimization. However, we can rule out clonal variation in the SARM1 mutant since experiments were conducted with a cell pool as described above in order to avoid clonal variation of single clones.

The mechanism behind biosynthesis of NAADP is still debated. CD38 was the first enzyme discovered to possess the ability of producing NAADP. However, it requires acidic pH to produce NAADP[11] -which does not match the characteristics of a cytosolic NAADP producer. HeLa cells do not express CD38 and hence, it is not surprising that inhibition of CD38 had no effect on S. aureus invasion in HeLa cells. However, NAADP production by HeLa cells was observed in absence of CD38[12]. Thus CD38independent NAADP generation is likely. SARM1 can produce NAADP at neutral pH[13] and is expressed in HeLa, thus providing a more promising candidate.

We agree with the reviewer that the reduction of S. aureus internalization after ablation of SARM1 is less pronounced than in other experiments of ours. This may be explained by NAADP originating from other enzymes, such as the recently discovered DUOX1, DUOX2, NOX1 and NOX2[14], which – with exception of DUOX2- possess a low expression even in HeLa cells. We add this to the discussion in the revised manuscript (line 579).

We can, however, rule out clonal variation for the inhibitory effect. As stated above we generated K.O. cell pools specifically to avoid inherent problems of clonality. Thus, we also detect some residual wildtype cells within our cell pools.

The involvement of lysosomal secretion is, again, predicated largely on the basis of pharmacological evidence. No direct evidence is provided for the insertion of lysosomal components into the plasma membrane, or for the release of lysosomal contents to the medium. Instead, inhibition of lysosomal exocytosis by vacuolin-1 is the sole source of evidence. However, vacuolin-1 is by no means a specific inhibitor of lysosomal secretion: it is now known to act primarily as a PIKfyve inhibitor and to cause massive distortion of the endocytic compartment, including gross swelling of endolysosomes. The modest (20-25%) inhibition observed when using synaptotagmin 7 knockout cells is similarly not convincing proof of the requirement for lysosomal secretion.

We agree with the reviewer that the manuscript will benefit from a functional analysis of lysosomal exocytosis and therefore conducted assays to investigate exocytosis in the revised manuscript. We previously showed i) by addition of specific antisera that LAMP1 transiently is exposed on the plasma membrane during ionomycin and pore-forming toxin challenge and ii) demonstrated the release of ASM activity into the culture medium under these conditions.[9] However, both measurements are not compatible with S. aureus infection, since LAMP1 antibodies also are non-specifically bound by protein A and another IgG-binding proteins on the S. aureus surface, which would bias the results. Since protein A also may serve as an adhesin in the investigated pathway, we cannot simply delete the ORF without changing other aspects of staphylococcal virulence. Further, FBS contains a ASM background activity that impedes activity measurements of cell culture medium. We previously removed this background activity by a specific heat-inactivation protocol.[9] However, S. aureus invasion is strongly reduced in culture medium containing this heat-inactivated FBS.

We therefore developed a luminescence assay based on split NanoLuc luciferase that enables detection of LAMP1 exposed on the plasma membrane without usage of antibodies (Figure 2, A-C). We added a section on the assay in the revised manuscript. Briefly, we generated reporter cells by fusing a short peptide fragment of NanoLuc called HiBiT between the signal peptide and the mature luminal domain of LAMP1 and stably expressed the resulting protein in HeLa cells by lentiviral transduction. The LgBiT protein domain of NanoLuc luciferase (Promega) as well as the substrate Furimazine are added to the culture medium. HiBiT can reconstitute a functional NanoLuc with LgBiT and process Furimazine when lysosomes are exocytosed thereby generating luminescence measurable in a suitable plate reader.

With this assay we detected that about 30% of lysosomes that were “releasable” by treatment with ionomycin are exocytosed during S. aureus infection. Lysosomal exocytosis was strongly reduced (even below the levels of untreated controls), if we treated cells with Vacuolin-1 or Ned19.

We agree with the reviewer that Vacuolin-1 to some extent has unspecific side effects as has been shown by others and which we addressed in the revised version of the manuscript (line 541 ff). However, our new results with the HiBiT reporter cell line clearly demonstrate a reduction of lysosomal exocytosis after Vacuolin-1 treatment. Supported by this and our other results we hypothesize that Vacuolin-1 decreases S. aureus internalization due to the inhibition of lysosomal exocytosis.

As to the involvement of synaptotagmin 7: The effect of Syt7 K.O. on invasion was moderate in initial experiments, likely due to a high culture passage and presumably overgrowth of WT cells. However, reduction of invasion in Syt7 K.O.s was more pronounced in experiments with β-toxin complementation (Figure 2, N) and hence, we combined the two data sets (Figure 2, F). This demonstrates the reduction of bacterial invasion by ~40% in Syt7 K.O. cell pools. Moreover, Syt7 is not the only protein possibly involved in Ca2+-dependent exocytosis. For instance, Syt1 has been shown to possess an overlapping function.[15] This may explain the differences between our Vacuolin-1 and Syt7 ablation experiments. We added this information to the discussion.

ASM is proposed to play a central role in the rapid invasion process. As above, most of the evidence offered in this regard is pharmacological and often inconsistent between inhibitors or among cell types. Some drugs affect some of the cells, but not others. It is difficult to reach general conclusions regarding the role of ASM. The argument is made even more complex by the authors' use of exogenous sphingomyelinase (beta-toxin). Pretreatment with the toxin decreased invasion efficiency, a seemingly paradoxical result. Incidentally, the effectiveness of the added toxin is never quantified/validated by directly measuring the generation of ceramide or the disappearance of SM.

Although pharmacological inhibitors can have unspecific side effects, we want to emphasize that the inhibitors used in our study act on the enzyme ASM by completely different mechanisms. Amitriptyline is a so called functional inhibitor of ASM (FIASMA) which induces the detachment of ASM from lysosomal membranes resulting in degradation of the enzyme.[16] By contrast, ARC39 is a competitive inhibitor.[17, 18]

There are no inconsistencies in our data obtained with ASM inhibitors. Amitriptyline and ARC39 both reduce the invasion of S. aureus in HuLEC, HuVEC and HeLa cells (Figure 2G). ARC39 needs a longer pre-incubation, since its uptake by host cells is slower (to be published elsewhere). We observe a different outcome in 16HBE14o- and Ea.Hy 926 cells, with 16HBE14o- even demonstrating a slightly increased invasion of S. aureus upon ARC39 treatment. Amitriptyline had no effect (Figure 2G).

Thus, the ASM-dependent S. aureus internalization is cell type/line specific, which we state in the manuscript. The molecular origin of these differences is unclear and will require further investigation, e.g. in testing cell lines for potential differences in surface receptors. In a separate study we have already developed a biotinylation-based approach to identify potential novel host cell surface interaction partners during S. aureus infection.[19]

Moreover, both inhibitors affected the invasion dynamics (Figure 3D), phagosomal escape (Figure 4C and Figure 4D) and Rab7 recruitment (Figure 4A and Supp. Figure 4A-C) in a similar fashion. Proper inhibition of ASM by both compounds in all cell lines used was validated by enzyme assays (Supp. Figure 2H), which again suggests that the ASM-dependent pathway does only exist in specific cell lines and also supports that we do not observe unspecific side effects of the compounds. We clarified this in the revised manuscript.

ASM is a key player for SM degradation and recycling. In clinical context, deficiency in ASM results in the so-called Niemann Pick disease type A/B. The lipid profile of ASM-deficient cells is massively altered[20], which will result in severe side effects. Short-term inhibition by small molecules therefore poses a clear benefit when compared to the usage of ASM K.O. cells. In order to satisfy the query of the reviewer, we generated two ASM K.O. cell pools (generated with two different sgRNAs) and tested these for S. aureus invasion efficiency (Figure 2, I). We did not observe bacterial invasion differences between WT and K.O. cells. However, when we treated the cells additionally with ASM inhibitor, we observed a strongly reduced invasion in WT cells, while invasion efficiency in ASM K.O. was only slightly affected (Figure 2, J). We concluded that the reduced invasion observed in inhibitor-treated WT cells predominantly is due to absence of ASM, while the small reduction observed in ARC39treated ASM K.O.s is likely due to unspecific side effects.

We performed lipidomics on these cells and demonstrated a strongly altered sphingolipid profile in ASM K.O. cells compared to untreated and inhibitor-treated WT cells (Figure 2, K). We speculate that other ASM-independent bacterial invasion pathways are upregulated in ASM K.O.s., thereby obscuring the effect contributed by absence of ASM. We discussed this in the revised manuscript (line 518 ff).

Moreover, we introduced the RFP-CWT escape marker into the ASM K.O. cells and measured phagosomal escape of S. aureus JE2 and Cowan I. The latter strain is non-cytotoxic and serves as negative control, since it is known to possess a very low escape rate, due to its inability to produce toxin. Again, we compared early invaders (infection for 10 min) with early+late invaders (infection for 30 min). As observed for JE2, “early invaders” possess lower escape rates than “early+late invaders”.

We did not observe differences between WT and ASM K.O. cells, if we infected for only 10 min. By contrast, we observed a lower escape rate in ASM K.O (Author response image 3, see end of the document). compared to WT cells, when we infected for 30 min.

However, we usually observe an increased phagosomal escape, when we treated host cells with ASM inhibitors (Figure 4C and D). Reduced phagosomal escape of intracellular S. aureus in ASM K.O. cells may be caused by the altered sphingolipid profile(e.g., by interference with binding of bacterial toxins to phagosomal membranes or altered vesicular acidification). We hence think that these data are difficult to interpret, and clarification would require intense additional experimentation. Thus, we did not include this data in the manuscript.

Phagosomal escape rates were established in either HeLa wild-type or ASM K.

O. cells expressing the phagosomal escape reporter RFP-CWT. Host cells that were infected with the cytotoxic S. aureus strain JE2 or the non-cytotoxic strain Cowan I for 10 or 30 minutes and escape rates were determined by microscopy 3h p.i.

As to the treatment with a bacterial sphingomyelinase:

Treatment with the bacterial SMase (bSMase, here: β-toxin) was performed in two different ways:

i) Pretreatment of host cells with β-toxin to remove SM from the host cell surface before infection. This removes the substrate of ASM from the cell surface prior to addition of the bacteria (Figure 2L, Figure 4A-C). Since SM is not present on the extracellular plasma membrane leaflet after treatment, a release of ASM cannot cause localized ceramide formation at the sites of lysosomal exocytosis. Similar observations were made by others.[21]

ii) Addition of bSMase to host cells together with the bacteria to complement for the absence of ASM (Figure 2N).

Removal of the ASM substrate before infection (i) prevents localized ASM-mediated conversion of SM to Cer during infection and resulted in a decreased invasion, while addition of the SMase during infection resulted in an increased invasion in TPC1 and Syt7 ablated cells. Thus, both experiments are consistent with each other and in line with our other observations.

Removal of SM from the plasma membrane by β-toxin was indirectly demonstrated by the absence of Lysenin recruitment to phagosomes/escaped bacteria when host cells were pretreatment with the toxin before infection (Figure5C). We also added another data set that demonstrates degradation of a fluorescence SM derivative upon β-toxin treatment of host cells (Supp Figure 2, M). In another publication, we recently quantified the effectiveness of β-toxin treatment, even though with slightly longer treatment times (75 min vs. 3h).[22]

To clarify our experimental approaches to the readership we added an explanatory section to the revised manuscript (line 287 ff) and we also added a scheme to in Figure 2M describing the experimental settings.

As to the general conclusions regarding the role of ASM: ASM and lysosomal exocytosis has been shown to be involved in uptake of a variety of pathogens[21, 23-27] supporting its role in the process.

The use of fluorescent analogs of sphingomyelin and ceramide is not well justified and it is unclear what conclusions can be derived from these observations. Despite the low resolution of the images provided, it appears as if the labeled lipids are largely in endomembrane compartments, where they would presumably be inaccessible to the secreted ASM. Moreover, considering the location of the BODIPY probe, the authors would be unable to distinguish intact sphingomyelin from its breakdown product, ceramide. What can be concluded from these experiments? Incidentally, the authors report only 10% of BODIPY-positive events after 10 min. What are the implications of this finding? That 90% of the invasion events are unrelated to sphingomyelin, ASM, and ceramide?

During the experiments with fluorescent SM analogues (Figure 3a,b), S. aureus was added to the samples immediately before the start of video recording. Hence, bacteria are slowly trickling onto the host cells, and we thus can image the initial contact between them and the bacteria, for instance, the bacteria depicted in Figure 3A contact the host cell about 9 min before becoming BODIPY-FL-positive (see Supp. Video 1, 55 min). Hence, in these cases we see the formation of phagosomes around bacteria rather than bacteria in endomembrane compartments. Since generation of phagosomes happens at the plasma membrane, SM is accessible to secreted ASM.

The “trickling” approach for infection is an experimental difference to our invasion measurements, in which we synchronized the infection by centrifugation. This ensures that all bacteria have contact to host cells and are not just floating in the culture medium. However, live cell imaging of initial bacterialhost contact and synchronization of infection is hard to combine technically.

In our invasion measurements -with synchronization-, we typically see internalization of ~20% of all added bacteria after 30 min. Hence, most bacteria that are visible in our videos likely are still extracellular and only a small proportion was internalized. This explains why only 10% of total bacteria are positive for BODIPY-FL-SM after 10 min. The proportion of internalized bacteria that are positive for BODIPY-FL-SM should be way higher but cannot be determined with this method.

We agree with the reviewer that we cannot observe conversion of BODIPY-FL-SM by ASM. In order to do that, we attempted to visualize the conversion of a visible-range SM FRET probe (Supp. Figure 3), but the structure of the probe is not compatible with measurement of conversion on the plasma membrane, since the FITC fluorophore released into the culture medium by the ASM activity thereby gets lost for imaging. In general, the visualization of SM conversion with subcellular resolution is challenging and even with novel tools developed in our lab[28] visualization of SM on the plasma membrane is difficult.

The conclusions we draw from these experiments are that i.) S. aureus invasion is associated with SM and ii.) SM-associated invasion can be very fast, since bacteria are rapidly engulfed by BODIPY-FL-SM containing membranes.

It is also unclear how the authors can distinguish lysenin entry into ruptured vacuoles from the entry of RFP-CWT, used as a criterion of bacterial escape. Surely the molecular weights of the probes are not sufficiently different to prevent the latter one from traversing the permeabilized membrane until such time that the bacteria escape from the vacuole.

We here want to clarify that both Lysenin as well as the CWT reporter have access to ruptured vacuoles (Figure 4B). We used the Lysenin reporter in these experiments for estimation of SM content of phagosomal membranes. If a vacuole is ruptured, both the bacteria and the luminal leaflet of the phagosomal membrane remnants get in contact with the cytosol and hence with the cytosolically expressed reporters YFP-Lysenin as well as RFP-CWT resulting in “Lysenin-positive escape” when phagosomes contained SM (see Figure 5C). By contrast, either β-toxin expression by S. aureus or pretreatment with the bSMase resulted in absence of Lysenin recruitment suggesting that the phagosomal SM levels were decreased/undetectable (Figure 5C, Supp Figure 6F, G, I, J).

Although this approach does not enable a quantitative measurement of phagosomal SM, this method is sufficient to show that β-toxin expression and pretreatment result in markedly decreased phagosomal SM levels in the host cells.

The approach we used here to analyze “Lysenin-positive escape” can clearly be distinguished from Lysenin-based methods that were used by others.29 There Lysenin was used to show trans-bilayer movement of SM before rupture of bacteria-containing phagosomes.

To clarify the function of Lysenin in our approach we added additional figures (Figure 4F, Supp. Figure 5) and a movie (Supp. Video 4) to the revised manuscript.

Both SMase inhibitors (Figure 4C) and SMase pretreatment increased bacterial escape from the vacuole. The former should prevent SM hydrolysis and formation of ceramide, while the latter treatment should have the exact opposite effects, yet the end result is the same. What can one conclude regarding the need and role of the SMase products in the escape process?

As pointed out above, pretreatment of host cells with SMase removes SM from the plasma membrane and hence, ASM does not have access to its substrate. Hence, both treatment with either ASM inhibitors or pretreatment with bacterial SMase prevent ASM from being active on the plasma membrane and hence block the ASM-dependent uptake (Figure 2 G, L). Although overall less bacteria were internalized by host cells under these conditions, the bacteria that invaded host cells did so in an ASM-independent manner.

Since blockage of the ASM-dependent internalization pathway (with ASM inhibitor [Figure 4C, D], SMase pretreatment [Figure 5B] and Vacuolin-1[Figure.4E]) always resulted in enhanced phagosomal escape, we conclude that bacteria that were internalized in an ASM-independent fashion cause enhanced escape. Vice versa, bacteria that enter host cells in an ASM-dependent manner demonstrate lower escape rates.

This is supported by comparing the escape rates of “early” and “late” invaders [Figure 5D, E], which in our opinion is a key experiment that supports this hypothesis. The “early” invaders are predominantly ASM-dependent (see e.g. Figure 3E) and thus, bacteria that entered host cell in the first 10 min of infection should have been internalized predominantly in an ASM-dependent fashion, while slower entry pathways are active later during infection. The early ASM dependent invaders possessed lower escape rates, which is in line with the data obtained with inhibitors (e.g. Figure 4C, D).

We hypothesize that the activity of ASM on the plasma membrane during invasion mediates the recruitment of a specific subset of receptors, which then influences downstream phagosomal maturation and escape. This hypothesis is supported by the fact that the subset of receptors interacting with S. aureus is altered upon inhibition of the ASM-dependent uptake pathway. We describe this in another study that is currently under evaluation elsewhere.

Reviewer #2 (Public review):

Summary:

In this manuscript, Ruhling et al propose a rapid uptake pathway that is dependent on lysosomal exocytosis, lysosomal Ca2+ and acid sphingomyelinase, and further suggest that the intracellular trafficking and fate of the pathogen is dictated by the mode of entry.

The evidence provided is solid, methods used are appropriate and results largely support their conclusions, but can be substantiated further as detailed below. The weakness is a reliance on chemical inhibitors that can be non-specific to delineate critical steps.

Specific comments:

A large number of experiments rely on treatment with chemical inhibitors. While this approach is reasonable, many of the inhibitors employed such as amitriptyline and vacuolin1 have other or nondefined cellular targets and pleiotropic effects cannot be ruled out. Given the centrality of ASM for the manuscript, it will be important to replicate some key results with ASM KO cells.

We thank the reviewer for the critical evaluation of our manuscript and plenty of constructive comments.

We agree with the reviewer, that ASM inhibitors such as functional inhibitors of ASM (FIASMA) like amitriptyline used in our study have unspecific side effects given their mode-of-action. FIASMAs induce the detachment of ASM from lysosomal membranes resulting in degradation of the enzyme.[16] However, we want to emphasize that we also used the competitive inhibitor ARC39 in our study[17, 18] which acts on the enzyme by a completely different mechanism. All phenotypes (reduced invasion [Figure 2G], effect on invasion dynamics [Figure 3D], enhanced escape [Figure 4C, D] and differential recruitment of Rab7 [Supp. Figure 4A-C]) were observed with both inhibitors thereby supporting the role of ASM in the process.

We further agree that experiments with genetic evidence usually support and improve scientific findings. However, ASM is a cellular key player for SM degradation and recycling. In a clinical context, deficiency in ASM results in a so-called Niemann Pick disease type A/B. The lipid profile of ASMdeficient cells is massively altered[20], which in itself will result in severe side effects. Thus, the usage of inhibitors provides a clear benefit when compared to ASM K.O. cells, since ASM activity can be targeted in a short-term fashion thereby preventing larger alterations in cellular lipid composition.

We nevertheless generated two ASM K.O. cell pools (generated with two different sgRNAs) and tested for invasion efficiency (Figure 2, I). Here, we did not observe differences between WT and mutants. However, if we treated the cells additionally with ASM inhibitor, we observed a strongly reduced invasion in WT cells, while invasion efficiency in ASM K.O. was only slightly affected (Figure 2, J). We concluded that the reduced invasion observed in WT cells upon inhibitor treatment predominantly is due to inhibition of ASM, whereas the small reduction observed in ARC39-treated ASM K.O.s is likely due to unspecific side effects. We also demonstrated a strongly altered sphingolipid profile in ASM K.O. cells when compared to untreated and inhibitor-treated WT cells (new Figure 2, K). We speculate that other ASM-independent invasion pathways are upregulated in ASM K.O.s., thereby making up for the absence of ASM. We discuss this in the revised manuscript (line 518 ff).

We introduced the RFP-CWT escape marker into the ASM K.O. cells and measured phagosomal escape of S. aureus JE2 and Cowan I (Author response image 3). The latter serves as negative control, since it is known to possess a very low escape rate, due to its inability of toxin production. Again, we compared early invaders (infection for 10 min) with early+late invaders (infection for 30 min). As seen before for JE2, early invaders possess lower escape rates than early+late invaders. We did not observe differences between WT and K.O. cells, if we infected for 10 min. By contrast, we observed a lower escape rate in ASM K.O. compared to WT cells, when we infected for 30 min. However, we usually observe an increased phagosomal escape, when we treated host cells with ASM inhibitors (Figure 4C and D). We think that the reduced phagosomal escape in ASM K.O. is caused by the altered sphingolipid profile, which could have versatile effects (e.g., inference with binding of bacterial toxins to phagosomal membranes or changes in acidification). We hence think that these data are difficult to interpret, and clarification would require intense additional experimentation. Thus, we did not include this data in the manuscript.

Most experiments are done in HeLa cells. Given the pathway is projected as generic, it will be important to further characterize cell type specificity for the process. Some evidence for a similar mechanism in other cell types S. aureus infects, perhaps phagocytic cell type, might be good.

Whenever possible we performed the experiments not only in HeLa but also in HuLECs. For example, we refer to experiments concerning the role of Ca2+ (Figure 1A/Supp.Figure1A), lysosomal Ca2+/Ned19 (Figure1B/Supp Figure 1C), lysosomal exocytosis/Vacuolin-1 (Figure 2D/Supp. Figure2D), ASM/ARC39 and amitriptyline (Figure 2G), surface SM/β-toxin (Figure 2L/Supp. Figure 2L), analysis of invasion dynamics (complete Figure 3) and measurement of cell death during infection (Figure 6C+E, Supp. Figure 8A+B).

HuLECs, however, are not really genetically amenable and hence we were not able to generate gene deletions in these cells and upon introduction of the fluorescence escape reporter the cells are not readily growing.

As to ASM involvement in phagocytic cells: a role for ASM during the uptake of S. aureus by macrophages was previously reported by others.[25] However, in professional phagocytes S. aureus does not escape from the phagosome and replicates within the phagosome.[30]

I'm a little confused about the role of ASM on the surface. Presumably, it converts SM to ceramide, as the final model suggests. Overexpression of b-toxin results in the near complete absence of SM on phagosomes (having representative images will help appreciate this), but why is phagosomal SM detected at high levels in untreated conditions? If bacteria are engulfed by SM-containing membrane compartments, what role does ASM play on the surface? If surface SM is necessary for phagosomal escape within the cell, do the authors imply that ASM is tuning the surface SM levels to a certain optimal range? Alternatively, can there be additional roles for ASM on the cell surface? Can surface SM levels be visualized (for example, in Figure 4 E, F)?

We initially hypothesized that we would detect higher phagosomal SM levels upon inhibition of ASM, since our model suggests SM cleavage by ASM on the host cell surface during bacterial cell entry. However, we did not detect any changes in our experiments (Supp. Figure 4F). We currently favor the following explanation: SM is the most abundant sphingolipid in human cells.[31] If peripheral lysosomes are exocytosed and thereby release ASM, only a localized and relative small proportion of SM may get converted to Cer, which most likely is below our detection limit. In addition, the detection of cytosolically exposed phagosomal SM by YFP-Lysenin is not quantitative and provides a “Yes or No” measurement. Hence, we think that the rather limited SM to Cer conversion in combination with the high abundance of SM in cellular membranes does not visibly affect the recruitment of the Lysenin reporter.

In our experiments that employ BODIPY-FL-SM (Figure 3a+b), we cannot distinguish between native SM and downstream metabolites such as Cer. Hence, again we cannot make any assumptions on the extent to which SM is converted on the surface during bacterial internalization. Although our laboratory recently used trifunctional sphingolipid analogs to analyze the SM to Cer conversion[22], the visualization of this process on the plasma membrane is currently still challenging.

Overall, we hypothesize that the localized generation of Cer on the surface by released ASM leads to generation of Cer-enriched platforms. Subsequently, a certain subset of receptors may be recruited to these platforms and influence the uptake process. These platforms are supposed to be very small, which also would explain that we did not detect changes in Lysenin recruitment.

Related to that, why is ASM activity on the cell surface important? Its role in non-infectious or other contexts can be discussed.

ASM release by lysosomal exocytosis is implied in plasma membrane repair upon injury. We added a short description of the role of extracellular ASM in the introduction (line 35).

If SM removal is so crucial for uptake, can exocytosis of lysosomes alone provide sufficient ASM for SM removal? How much or to what extent is lysosomal exocytosis enhanced by initial signaling events? Do the authors envisage the early events in their model happening in localized confines of the PM, this can be discussed.

Ionomycin treatment led to a release of ~10 % of all lysosomes and also increased extracellular ASM activity.[8, 9] In the revised manuscript, we developed an assay to determine lysosomal exocytosis during S. aureus infection (Figure 2, A-C). We detected lysosomal exocytosis of ~30% when compared to ionomycin treatment during infection. Since this is only a fraction of the “releasable lysosomes”, we assume that the effects (lysosomal Ca2+ liberation, lysosomal exocytosis and ASM activity) are very localized and take place only at host-pathogen contact sites (see also above). We discuss this in the revised manuscript (line 563 ff). To our knowledge it is currently unclear to which extent the released ASM affects surface SM levels. We attempted to visualize the local ASM activity on the cell surface by using a visible range FRET probe (Supp. Fig. 3). Cleavage of the probe by ASM on the surface leads to release of FITC into the cell culture medium, which does not contribute a measurable signal at the surface.

How are inhibitor doses determined? How efficient is the removal of extracellular bacteria at 10 min? It will be good to substantiate the cfu experiments for infectivity with imaging-based methods. Are the roles of TPC1 and TPC2 redundant? If so, why does silencing TPC1 alone result in a decrease in infectivity? For these and other assays, it would be better to show raw values for infectivity. Please show alterations in lysosomal Ca2+ at the doses of inhibitors indicated. Is lysosomal Ca2+ released upon S. aureus binding to the cell surface? Will be good to directly visualize this.

Concerning the inhibitor concentrations, we either used values established in published studies or recommendations of the suppliers (e.g. 2-APB, Ned19, Vacuolin-1). For ASM inhibitors, we determined proper inhibition of ASM by activity assays. Concentrations of ionomycin resulting in Ca2+ influx and lysosomal exocytosis was determined in earlier studies of our lab.[9, 32]

As to the removal of bacteria at 10 min p.i.: Lysostaphin is very efficient for removal of extracellular S. aureus and sterilizes the tissue culture supernatant. It significantly lyses bacteria within a few minutes, as determined by turbidity assays.[33]

As to imaging-based infectivity assays: We performed imaging-based invasion assays to show reduced invasion efficiency with two ASM inhibitors in the revised manuscript with similar results as obtained by CFU counts (Supp. Figure 2, J).

Regarding the roles of TPC1 and TPC2: from our data we cannot conclude whether the roles of TPC1 and TPC2 are redundant. One could speculate that since blockage of TPC1 alone is sufficient to reduce internalization of bacteria, that both channels may have distinct roles. On the other hand, there might be a Ca2+ threshold in order to initiate lysosomal exocytosis that can only be attained if TPC1 and TPC2 are activated in parallel. Thus, our observations are in line with another study that shows reduced Ebola virus infection in absence of either TPC1 or TPC2.[34] In order to address the role of TPC2 for this review process, we kindly were gifted TPCN1/TPCN2 double knock-out HeLa cells by Norbert Klugbauer (Freiburg, Germany), which we tested for S. aureus internalization. We found that invasion was reduced in these double KO cell lines even further supporting a role of lysosomal Ca2+ release in S. aureus host cell entry (Author response image 2, see end of the document). Since we did not have a single TPCN2 knockout available, we decided to exclude these data from the main manuscript.

As to raw CFU counts: whereas the observed effects upon blocking the invasion of S. aureus are stable, the number of internalized bacteria varies between individual biological replicates, for instance, by differences in host cell fitness or growth differences in bacterial cultures, which are prepared freshly for each experiment.

With respect to visualization of lysosomal Ca2+ release: we agree with the reviewer that direct visual demonstration of lysosomal Ca2+ release upon infection would improve the manuscript. We therefore performed live cell imaging to visualize lysosomal Ca2+ release by a previously published method.[1] The approach is based on two dextran-coupled fluorophores that were incubated with host cells. The dyes are endocytosed and eventually stain the lysosomes. One of the dyes, Rhod-2, is Ca2+-sensitive and can be used to estimate the lysosomal Ca2+ content. The second dye, AF647, is Ca2+-insensitive and is used to visualize the lysosomes. If the ratio Rhod-2/AF647 within the lysosomes is decreasing, lysosomal Ca2+ release is indicated. We monitored lysosomal Ca2+ content during S. aureus infection with this method (Author response image 1 and Author response video 1). However, the lysosomes are very dynamic, and it is challenging to monitor the fluorescence intensities over time. Thus, quantitative measurements are not possible with our methodology, and we decided to not include these data in the final manuscript. However, one could speculate that lysosomal Ca2+ content in the selected ROI (Author response image 1 and Author response video 1) is decreased upon attachment of S. aureus to the host cells as indicated by a decrease in Rhod-2/AF647 ratio.

The precise identification of cytosolic vs phagosomal bacteria is not very easy to appreciate. The methods section indicates how this distinction is made, but how do the authors deal with partial overlaps and ambiguities generally associated with such analyses? Please show respective images.

The number of events (individual bacteria) for the live cell imaging data should be clearly mentioned.

We apologize for not having sufficiently explained the technology to detect escaped S. aureus. The cytosolic location of S. aureus is indicated by recruitment of RFP-CWT.[35] CWT is the cell wall targeting domain of lysostaphin, which efficiently binds to the pentaglycine cross bridge in the peptidoglycan of S. aureus. This reporter is exclusively and homogenously expressed in the host cytosol. Only upon rupture of phagoendosomal membranes, the reporter can be recruited to the cell wall of now cytosolically located bacteria. S. aureus mutants, for instance in the agr quorum sensing system, cannot break down the phagosomal membrane in non-professional phagocytes and thus stay unlabeled by the CWT-reporter.[35] We include several images (Figure 4, F, Supp. Figure 5) /movies (Supp. Video 4) of escape events in the revised manuscript. The bacteria numbers for live cell experiments are now shown in Supp. Figure 7.

In the phagosome maturation experiments, what is the proportion of bacteria in Rab5 or Rab7 compartments at each time point? Will the decreased Rab7 association be accompanied by increased Rab5? Showing raw values and images will help appreciate such differences. Given the expertise and tools available in live cell imaging, can the authors trace Rab5 and Rab7 positive compartment times for the same bacteria?

We included the proportion of Rab7-associated bacteria in the revised manuscript (Supp. Figure 4A and C) and also shortly mention these proportions in the text (line 353). Usually, we observe that Rab5 is only transiently (for a few minutes) present on phagosomes and only afterwards the phagosomes become positive for Rab7. We do not think that a decrease in Rab7-positive phagosomes would increase the proportion of Rab5-positive phagosomes. However, we cannot exclude this hypothesis with our data.

We can achieve tracing of individual bacteria for recruitment of Rab5/Rab7 only manually, which impedes a quantitative evaluation. However, we included a Video (Supp. Video 3) that illustrates the consecutive recruitment of the GTPases.

The results with longer-term infection are interesting. Live cell imaging suggests that ASM-inhibited cells show accelerated phagosomal escape that reduces by 6 hpi. Where are the bacteria at this time point ? Presumably, they should have reached lysosomes. The relationship between cytosolic escape, replication, and host cell death is interesting, but the evidence, as presented is correlative for the populations. Given the use of live cell imaging, can the authors show these events in the same cell?

We think that most bacteria-containing phagoendosomes should have fused with lysosomes 6 h p.i. as we have previously shown by acidification to pH of 5 and LAMP1 decoration.[36]

The correlation between phagosomal escape and replication in the cytosol of non-professional phagocytes has been observed by us and others. In the revised manuscript we also provide images (Supp. Figure 5)/videos (Supp. Video 4) to show this correlation in our experiments.

Given the inherent heterogeneity in uptake processes and the use of inhibitors in most experiments, the distinction between ASM-dependent and independent pathways might not be as clear-cut as the authors suggest. Some caution here will be good. Can the authors estimate what fraction of intracellular bacteria are taken up ASM-dependent?

We agree with the reviewer that an overlap between internalization pathways is likely. A clear distinction is therefore certainly non-trivial. Alternative to ASM-dependent and ASM-independent pathways, the ASM activity may also accelerate one or several internalization pathways. We address this limitation in the discussion of the revised manuscript (line 596 ff).

Early in infection (~10 min after contact with the cells), the proportion of bacteria that enter host cells ASM-dependently is relatively high amounting to roughly 75-80% in HuLEC. After 30 min, this proportion is decreasing to about 50%. We included a paragraph in the discussion of the revised manuscript (line 593 ff).

Reviewer #2 (Recommendations for the authors):

(1) The experiment in Figure 4H is interesting. Details on what proportion of the cell is double positive, and if only this fraction was used for analysis will be good.

We did use all bacteria found in the images independently from whether host cells were infected with only one or both strains. We unfortunately cannot properly determine the proportion of cells that are double infected, since i) we record the samples with CLSM and hence, cannot exclude that there are intracellular bacteria found in higher or lower optical sections. ii) we visualized cells by staining Nuclei and did not stain the cell borders, thus we cannot precisely tell to which host cell the bacteria localize.

(2) Data is sparse for steps 5 and 6 of the model (line 330).

We apologize for the inconvenience. There is a related study published elsewhere[19], in which we identified NRCAM and PTK7 as putative receptors involved in this invasion pathway. We included a section in the discussion with the corresponding citation (line 569).

(3) Data for the reduced number of intracellular bacteria upon blocking ASM-dependent uptake (line 235) is not clear. Do they mean decreased invasion efficiency? These two need not be the same.

We changed “reduced number of intracellular bacteria” to “invasion efficiency”.

(4) b-toxin added to the surface can get endocytosed. Can its surface effect be delineated from endo/phagosomal effect?

We attempted to delineate effects contributed by the toxin activity on the surface vs. within phagosomes (Figure 5 A-C). We see an increased phagosomal escape, when we pretreated host cells with β-toxin (removal of SM form the surface) and infected either in presence (toxin will be taken up together with the bacteria into the phagosome) or in absence (toxin was washed away shortly before infection) of β-toxin. By contrast, overexpression of β-toxin by S. aureus did not affect phagosomal escape rates. The proper activity of β-toxin was confirmed by absence of Lysenin recruitment during phagosomal escape in all three conditions. We concluded that the activity on the surface and not the activity in the phagosome is important.

(5) The potential role(s) of bacterial factors in the uptake and subsequent intracellular stages can be discussed.

There are multiple bacterial adhesins known in S. aureus. These usually are either covalently attached to the bacterial cell wall such as the sortase-dependently anchored Fibronectin-binding Proteins A and B but also secreted and “cell wall binding” proteins as well at non proteinaceous factor such as wall-teichoic acids. A discussion of these factors would thus be out of the scope of this manuscript, and we here suggest reverting to specialized reviews on that topic.

(6) The manuscript is not very easy to read. The abstract could be rephrased for better clarity and succinctness, with a clearly stated problem statement. The introduction is somewhat haphazard, I feel it can be better structured.

We apologize for the inconvenience. We stated the problem/research question in the abstract and tried to improve the introduction without adding too much unnecessary detail. In general, we tried to improve the readability of the manuscript and hope that our results and conclusions can be easier understood by the reader in the revised version.

(7) Typo in Figure 5F. Step 6 should read "accessory receptors"

The typo was corrected.

References

(1) Lloyd-Evans, E. et al. Niemann-Pick disease type C1 is a sphingosine storage disease that causes deregulation of lysosomal calcium. Nature Medicine 14, 1247-1255 (2008).

(2) Launay, P. et al. TRPM4 Is a Ca2+-Activated Nonselective Cation Channel Mediating Cell Membrane Depolarization. Cell 109, 397-407 (2002).

(3) Nilius, B. et al. The Ca2+‐activated cation channel TRPM4 is regulated by phosphatidylinositol 4,5‐biphosphate. The EMBO Journal 25, 467-478-478 (2006).

(4) Cáceres, M. et al. TRPM4 Is a Novel Component of the Adhesome Required for Focal Adhesion Disassembly, Migration and Contractility. PLoS One 10, e0130540 (2015).

(5) Silva, I., Brunett, M., Cáceres, M. & Cerda, O. TRPM4 modulates focal adhesion-associated calcium signals and dynamics. Biophysical Journal 123, 390a (2024).

(6) Schlesier, T., Siegmund, A., Rescher, U. & Heilmann, C. Characterization of the Atl-mediated staphylococcal internalization mechanism. International Journal of Medical Microbiology 310, 151463 (2020).

(7) Jevon, M. et al. Mechanisms of Internalization ofStaphylococcus aureus by Cultured Human Osteoblasts. Infection and Immunity 67, 2677-2681 (1999).

(8) Rodriguez, A., Webster, P., Ortego, J. & Andrews, N.W. Lysosomes behave as Ca2+-regulated exocytic vesicles in fibroblasts and epithelial cells. J Cell Biol 137, 93-104 (1997).

(9) Krones & Rühling et al. Staphylococcus aureus alpha-Toxin Induces Acid Sphingomyelinase Release From a Human Endothelial Cell Line. Front Microbiol 12, 694489 (2021).

(10) Sakurai, Y. et al. Two-pore channels control Ebola virus host cell entry and are drug targets for disease treatment. Science 347, 995-998 (2015).

(11) Aarhus, R., Graeff, R.M., Dickey, D.M., Walseth, T.F. & Lee, H.C. ADP-ribosyl cyclase and CD38 catalyze the synthesis of a calcium-mobilizing metabolite from NADP. J Biol Chem 270, 3032730333 (1995).

(12) Schmid, F., Fliegert, R., Westphal, T., Bauche, A. & Guse, A.H. Nicotinic acid adenine dinucleotide phosphate (NAADP) degradation by alkaline phosphatase. J Biol Chem 287, 32525-32534 (2012).

(13) Angeletti, C. et al. SARM1 is a multi-functional NAD(P)ase with prominent base exchange activity, all regulated bymultiple physiologically relevant NAD metabolites. iScience 25, 103812 (2022).

(14) Gu, F. et al. Dual NADPH oxidases DUOX1 and DUOX2 synthesize NAADP and are necessary for Ca(2+) signaling during T cell activation. Sci Signal 14, eabe3800 (2021).

(15) Schonn, J.-S., Maximov, A., Lao, Y., Südhof, T.C. & Sørensen, J.B. Synaptotagmin-1 and -7 are functionally overlapping Ca2+ sensors for exocytosis in adrenal chromaffin cells. Proceedings of the National Academy of Sciences 105, 3998-4003 (2008).

(16) Kornhuber, J. et al. Functional Inhibitors of Acid Sphingomyelinase (FIASMAs): a novel pharmacological group of drugs with broad clinical applications. Cell Physiol Biochem 26, 9-20 (2010).

(17) Naser, E. et al. Characterization of the small molecule ARC39, a direct and specific inhibitor of acid sphingomyelinase in vitro. J Lipid Res 61, 896-910 (2020).

(18) Roth, A.G. et al. Potent and selective inhibition of acid sphingomyelinase by bisphosphonates. Angew Chem Int Ed Engl 48, 7560-7563 (2009).

(19) Rühling, M., Schmelz, F., Kempf, A., Paprotka, K. & Fraunholz Martin, J. Identification of the Staphylococcus aureus endothelial cell surface interactome by proximity labeling. mBio 0, e03654-03624 (2025).

(20) Schuchman, E.H. & Desnick, R.J. Types A and B Niemann-Pick disease. Mol Genet Metab 120, 27-33 (2017).

(21) Miller, M.E., Adhikary, S., Kolokoltsov, A.A. & Davey, R.A. Ebolavirus Requires Acid Sphingomyelinase Activity and Plasma Membrane Sphingomyelin for Infection. Journal of Virology 86, 7473-7483 (2012).

(22) M. Rühling, L.K., F. Wagner, F. Schumacher, D. Wigger, D. A. Helmerich, T. Pfeuffer, R. Elflein, C. Kappe, M. Sauer, C. Arenz, B. Kleuser, T. Rudel, M. Fraunholz, J. Seibel Trifunctional sphingomyelin derivatives enable nanoscale resolution of sphingomyelin turnover in physiological and infection processes via expansion microscopy. Nat Commun accepted in principle (2024).

(23) Peters, S. et al. Neisseria meningitidis Type IV Pili Trigger Ca(2+)-Dependent Lysosomal Trafficking of the Acid Sphingomyelinase To Enhance Surface Ceramide Levels. Infect Immun 87 (2019).

(24) Grassmé, H. et al. Acidic sphingomyelinase mediates entry of N. gonorrhoeae into nonphagocytic cells. Cell 91, 605-615 (1997).

(25) Li, C. et al. Regulation of Staphylococcus aureus Infection of Macrophages by CD44, Reactive Oxygen Species, and Acid Sphingomyelinase. Antioxid Redox Signal 28, 916-934 (2018).

(26) Fernandes, M.C. et al. Trypanosoma cruzi subverts the sphingomyelinase-mediated plasma membrane repair pathway for cell invasion. J Exp Med 208, 909-921 (2011).

(27) Luisoni, S. et al. Co-option of Membrane Wounding Enables Virus Penetration into Cells. Cell Host & Microbe 18, 75-85 (2015).

(28) Rühling, M. et al. Trifunctional sphingomyelin derivatives enable nanoscale resolution of sphingomyelin turnover in physiological and infection processes via expansion microscopy. Nature Communications 15, 7456 (2024).

(29) Ellison, C.J., Kukulski, W., Boyle, K.B., Munro, S. & Randow, F. Transbilayer Movement of Sphingomyelin Precedes Catastrophic Breakage of Enterobacteria-Containing Vacuoles. Curr Biol 30, 2974-2983 e2976 (2020).

(30) Moldovan, A. & Fraunholz, M.J. In or out: Phagosomal escape of Staphylococcus aureus. Cell Microbiol 21, e12997 (2019).

(31) Slotte, J.P. Biological functions of sphingomyelins. Progress in Lipid Research 52, 424-437 (2013).

(32) Stelzner, K. et al. Intracellular Staphylococcus aureus Perturbs the Host Cell Ca(2+) Homeostasis To Promote Cell Death. mBio 11 (2020).

(33) Kunz, T.C. et al. The Expandables: Cracking the Staphylococcal Cell Wall for Expansion Microscopy. Front Cell Infect Microbiol 11, 644750 (2021).

(34) Sakurai, Y. et al. Ebola virus. Two-pore channels control Ebola virus host cell entry and are drug targets for disease treatment. Science 347, 995-998 (2015).

(35) Grosz, M. et al. Cytoplasmic replication of Staphylococcus aureus upon phagosomal escape triggered by phenol-soluble modulin alpha. Cell Microbiol 16, 451-465 (2014).

(36) Giese, B. et al. Staphylococcal alpha-toxin is not sufficient to mediate escape from phagolysosomes in upper-airway epithelial cells. Infect Immun 77, 3611-3625 (2009).

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Staphylococcus aureus is an opportunistic pathogen causing severe diseases. Recently, S. aureus was recognized as intracellular pathogen with the intracellular niche promoting immune evasion and antibiotic resistance. We identified an alternative mechanism governing cellular uptake of S. aureus which relies on lysosomal Ca2+, lysosomal exocytosis and occurs concurrently to other well-known entry pathways within the same host cell population. This internalization pathway is rapid and active within only few minutes after bacterial contact with host cells. Compared to slow bacterial internalization, the rapid pathway demonstrates altered phagosomal maturation as well as translocation of the pathogen to the host cytosol and ultimately results in different rates of intracellular bacterial replication and host cell death. We show that these alternative infection outcomes are caused by the mode of bacterial uptake.

Revised version: Staphylococcus aureus is an opportunistic pathogen causing severe diseases. Recently, S. aureus was recognized as intracellular pathogen, whereby the intracellular niche promotes immune evasion and antibiotic resistance. Interaction of S. aureus with versatile host cell receptors was described previously, suggesting that internalization of the pathogen can occur via several pathways. It remains elusive whether the pathway of internalization can affect the intracellular fate of the bacteria. Here, we identified a mechanism governing cellular uptake of S. aureus which relies on lysosomal Ca2+, lysosomal exocytosis and occurs concurrently to other well-known entry pathways within the same host cell population. This internalization pathway is rapid and active within only few minutes after bacterial contact with host cells. Compared to slow bacterial internalization, the rapid pathway demonstrates altered phagosomal maturation as well as translocation of the pathogen to the host cytosol and ultimately results in different rates of intracellular bacterial replication and host cell death. We show that these alternative infection outcomes are caused by the mode of bacterial uptake.

Introduction

Earlier version: The most abundant sphingolipid, sphingomyelin (SM), resides in the extracellular leaflet of the plasma membrane and hence is exposed to the environment (1, 2). SM is a substrate for sphingomyelinases (SMases) that cleave off the phosphocholine head group to produce the non-polar ceramide. Acid sphingomyelinase (ASM) is a lysosomal enzyme involved in recycling of SM (3, 4) and is linked to infectious diseases, where it was shown to be involved in cellular uptake of several human pathogens (5–11).

Revised version: Sphingolipids are important components of eukaryotic membranes and also serve as bioactive signaling molecules (1). The most abundant sphingolipid, sphingomyelin (SM), resides in the extracellular leaflet of the plasma membrane and hence is exposed to the environment (2, 3). SM is a substrate for sphingomyelinases (SMases) that cleave off the phosphocholine head group to produce ceramide. Acid sphingomyelinase (ASM) is a lysosomal enzyme involved in recycling of SM (4, 5). Several studies demonstrated that ASM can be released by human cells either by Ca2+-dependent liberation of lysosomal content via so called lysosomal exocytosis (6-8) or by channeling the enzyme through the Golgi secretory pathway (9, 10). Extracellular ASM is active on the plasma membrane, where it has important roles during repair of membrane wounds (6, 11) but has also been shown to act on the SM-containing low-density lipoprotein (12). The release of ASM through lysosomal exocytosis mediates host cell entry of several human pathogens (7, 8, 13-17).

Earlier version: Staphylococcus aureus is a Gram-positive human commensal that asymptomatically colonizes about one third of the human population (12). It is an opportunistic pathogen causing diseases ranging from soft tissue and skin infections (13) to lethal diseases (14, 15). S. aureus is notorious for acquiring antibiotic resistances thereby causing >100,000 annual deaths (16). S. aureus possesses intracellular virulence strategies (17, 18) that contribute to antibiotic resistance (19, 20) and immune evasion (21). Internalization of S. aureus by host cells is mediated by several adhesins on the staphylococcal surface that bind a plethora of receptors on the host cell plasma membrane (5, 22–32). For instance, staphylococcal fibronectin-binding proteins that form fibronectin bridges to α5β1 integrins on host cells to initiate internalization (22) or clumping factor B which can interact with the host receptor annexin A2 (25).

Revised version: Staphylococcus aureus is a Gram-positive human commensal that asymptomatically colonizes about one third of the human population (18). It is an opportunistic pathogen causing diseases ranging from soft tissue and skin infections (19) to lethal diseases (20, 21). S. aureus is notorious for acquiring antibiotic resistances thereby causing >100,000 annual (22). The pathogen possesses intracellular virulence strategies (23, 24) which contribute to antibiotic resistance (25, 26) and immune evasion (27). Internalization of S. aureus by host cells is mediated by several adhesins on the staphylococcal surface that bind a plethora of receptors on the host cell plasma membrane (13, 28-38). For instance, staphylococcal fibronectin-binding proteins that form fibronectin bridges to α5β1 integrins on host cells to initiate internalization (28) or clumping factor B which can interact with the host receptor annexin A2 (31).

Earlier version: After internalization, S. aureus resides within a phagosome-like compartment that matures by acquiring proteins associated with early and then late endosomes as well as lysosomes (33–35). In epithelial and endothelial cells, the bacteria translocate to the host cytosol (“phagosomal escape”) where they replicate and cause host cell death (33, 36–38).

Revised version: After internalization, S. aureus resides within a phagosome-like compartment that matures by acquiring proteins associated with early and then late endosomes as well as lysosomes (39-41). In epithelial and endothelial cells, the bacteria translocate to the host cytosol ["phagosomal escape”, (39, 42, 43)]. Phagosomal escape is mediated by so-called phenol-soluble modulins (42), which comprise a family of helical amphiphilic peptides and are transcriptionally controlled by the accessory gene regulator agr, a staphylococcal quorum sensing system (44). S. aureus replicates in the host cytosol and cause host cell death for example by expression of a cysteine protease. (45).

Earlier version: The mechanism by which pathogens enter their host cells has been suggested to dictate the outcome of infections (39–42). For instance, phagosomes formed during natural internalization of Brucella arbortus by macrophages possessed different characteristics compared to phagosomes that were artificially generated by phorbol myristate acetate (42). The protozoic parasite Toxoplasma gondii invades its host cells by a proactive mechanism. Phagosomes generated during active invasion differed from those that were generated during Fc receptor-mediated uptake of antibody-coated parasites (41). Since previous observation connecting host cell entry and intracellular fate of pathogens were based on artificial induction of internalization, it is unclear whether pathogens can enter host cells within the same population via different naturally occurring mechanisms and if the entry route has a direct influence on the outcome of an intracellular infection.

Revised version: The mechanism of pathogen entry into host cells has been suggested to dictate the outcome of infections (46-49). For instance, phagosomes formed during natural internalization of Brucella abortus by macrophages exhibited different characteristics compared to phagosomes that were artificially generated by phorbol myristate acetate (49). The protozoan parasite Toxoplasma gondii actively invades its host cells. Phagosomes generated during this active invasion differed from those that were generated during Fc receptor-mediated uptake of antibody-coated parasites (48). Since previous observation connecting host cell entry and intracellular fate of pathogens were based on artificial induction of internalization, it is unclear whether pathogens can enter host cells within the same population via different naturally occurring mechanisms and if the entry route has a direct influence on the outcome of an intracellular infection.

Earlier version: Here, we describe the internalization of S. aureus by tissue cells via a rapid pathway taking place within minutes after contact between bacteria and host cell surface. This pathway requires nicotinic acid adenosine dinucleotide phosphate (NAADP)-dependent mobilization of lysosomal Ca2+, followed by lysosomal exocytosis and thereby release of ASM. Since the rapid uptake is concurrent to previously known internalization pathways, it probably has been missed due to long infection times used in traditional infection protocols (e.g., (22, 25, 29, 30)).

Revised version: Here, we describe the internalization of S. aureus by tissue cells via a rapid pathway taking place within minutes after contact between bacteria and host cell surface. This pathway requires nicotinic acid adenosine dinucleotide phosphate (NAADP)-dependent mobilization of lysosomal Ca2+, followed by lysosomal exocytosis and thereby release of ASM. Since the rapid uptake is concurrent to previously described internalization pathways, it probably has been missed due to long infection times used in traditional infection protocols [e.g., (28, 31, 35, 36)].

S. aureus bacteria, which enter host cells of the same cell population via the rapid pathway, cause a distinct infection outcome with altered phagosome maturation, bacterial translocation to host cytosol and eventually host cell death when compared to bacteria that enter host cells at later time points during infection. Thus, the outcome of an infection is decided at the single cell level during bacterial uptake at the host plasma membrane. Bacteria-containing phagosomes that were formed by the ASM-dependent rapid uptake pathway are delayed in their maturation and the bacteria do not escape these organelles efficiently, hence demonstrating a direct link between concurrently active modes of bacterial uptake and the resulting outcomes of intracellular S. aureus infection.

Results

S. aureus triggers lysosomal Ca2+ mobilization for host cell invasion

Added in the revised version: We previously showed that internalization of S. aureus by host cells is associated with cellular Ca2+ signaling (50). To validate these findings, we treated human microvascular endothelial cells (HuLEC, Figure 1, A) or HeLa (Supp. Figure 1, A) with the cell permeant Ca2+ chelator BAPTA-AM (51) to interfere with host cell Ca2+ signaling and subsequently infected the cells with S. aureus. We measured the number of internalized S. aureus and determined the invasion efficiency by normalizing bacteria numbers detected in treated samples to untreated controls (set to 100%). Invasion efficiency was reduced by BAPTA-AM in a concentration-dependent manner, indicating an involvement of Ca2+ in the host cell entry process of S. aureus.

Removed in the revised version: We previously identified that internalization of S. aureus by host cells is associated with cellular Ca2+ signaling (43). To delineate the role of Ca2+ in this process mechanistically, we pretreated human microvascular endothelial cells (HuLEC) with the calcium ionophore ionomycin, which reduced bacterial invasion efficiency dose-dependently (Figure 1, A), at concentrations neither affecting bacterial (Supp. Figure 1, A, B) nor host cell viability (Supp. Figure 1, C).

Earlier version: S. aureus invasion is dependent on Ca2+ liberation from lysosomal stores and lysosomal exocytosis.

Revised version: S. aureus invasion is dependent on Ca2+ liberation from lysosomal stores.

Earlier version: Treatment with Ionomycin (A, n≥3), 9-Phenantrol (B, n≥3), BAPTA-AM (C, n≥4), trans-Ned19 (D, n≥4), but not 2-APB and 8-Bromo-cADPR (D, n≥4) reduce S. aureus internalization by host cells in a TPC1 dependent fashion (E, n≥5). HuLEC (A, C, D, F), HeLa WT (B), HeLa TPC1 K.O. (E) or HeLa SARM1 K.O. (F) cells were treated with the respective substance and were subsequently infected with S. aureus JE2 for 30 min if not indicated otherwise. Extracellular bacteria were removed by lysostaphin, and the number of intracellular bacteria was determined by CFU counting. Results were normalized to untreated controls to obtain invasion efficiency (in percent of control). (F) scheme of host cell Ca2+ signaling and interfering agents. Statistics: one sample t-test (A-D, F), unpaired Student’s t-test (E). Bars represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001.

Revised version: Treatment with, BAPTA-AM (A, n≥4), trans-Ned19 (B, n≥4), but not 2-APB and 8-Bromo-cADPR (B, n≥4) reduce S. aureus internalization by host cells. Genetic ablation of TPC1 (C, n=5) and SARM1 (D, n=8) reduced invasion of S. aureus 10 min p.i. HuLEC (A, B), HeLa WT (C, D), HeLa TPC1 K.O. (C) or HeLa SARM1 K.O. (D) cells were treated with the respective substance and were subsequently infected with S. aureus JE2 for 30 min if not indicated otherwise. Extracellular bacteria were removed by lysostaphin, and the number of intracellular bacteria was determined by CFU counting. Results were normalized to untreated controls to obtain invasion efficiency (in percent of control). (F) Scheme of host cell Ca2+ signaling and interfering agents. Letters indicate figure panels that supported the respective conclusion. Statistics: one sample t-test (A, B,D,), unpaired Student’s t-test (C). Bars represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001. Created in BioRender.

Added in the revised version: The cytosolic Ca2+ levels, which are crucial for signaling processes, can be elevated by Ca2+ influx from the extracellular space. To test whether S. aureus internalization is mediated by Ca2+ influx, we measured invasion efficiency, in absence of Ca2+ from the cell culture medium (Figure 1, A and Supp. Figure 1, A). Internalization of S. aureus by host cells was not affected by absence of Ca2+ (“Ca2+-free, 0 µM BAPTA-AM”), thereby excluding a Ca2+ influx as crucial for S. aureus host cell entry. The pore-forming staphylococcal α-toxin was shown to mediate Ca2+ influx into host cells (11). S. aureus invasion into host cells was largely independent of the production of α-toxin (Supp. Figure 1, B), which further supported that the internalization is independent from extracellular Ca2+.

Removed in the revised version: Created with BioRender.com/k99z626.

Removed in the revised version: Since focal adhesions constitute major cellular entry sites for S. aureus and the Ca2+-activated transient receptor potential ion channel 4 (TRPM4) is located therein, we inhibited TRPM4 with 9-Phenantrol. 9-Phenantrol reduced invasion efficiency (Figure 1, B), but was not bactericidal (Supp. Figure 1, D). Further, treatment of HuLEC (Figure 1, C) or HeLa (Supp. Figure 1, E) with the Ca2+ chelator BAPTA-AM (44) led to reduction of S. aureus invasion in a concentration-dependent manner.

Removed in the revised version: The pore-forming staphylococcal α-toxin has been shown to mediate Ca2+ influx into host cells (45). However, S. aureus invasion was independent of toxin production in HeLa and only slightly reduced in endothelial cells (Supp. Figure 1, F). Consequently, omission of Ca2+ from the cell culture medium did not affect internalization of bacteria, indicating that invasion is not dependent on Ca2+ influx (Figure 1, C).

Earlier version: Since intracellular Ca2+ stores can serve as alternative source of Ca2+, we next blocked Ca2+ liberation from the ER using 2-APB (46) or 8-Bromo-cyclic ADP ribose (8-Bromo-cADPR; (47)) and from lysosomes using trans-Ned19 (48) in HuLEC (Figure 1, D) or HeLa (Supp. Figure 1, G, H) at concentrations not affecting bacterial viability (Supp. Figure 1, I and J). Interference with lysosomal, but not ER Ca2+ release reduced internalization of S. aureus.

Revised version: Since intracellular Ca2+ stores, such as the ER or lysosomes, can serve as alternative source of Ca2+ (52), we next blocked Ca2+ liberation either from the ER using 2-APB (53) or 8-Bromo-cyclic ADP ribose (8-Bromo-cADPR; (54)) or from lysosomes using trans-Ned19 (55). Interference with lysosomal but not ER Ca2+ release reduced internalization of S. aureus in HuLEC (Figure 1, B) or HeLa (Supp. Figure 1, C, D) at concentrations which do not affect bacterial viability (Supp. Figure 1, E and F).

Earlier version: Ned19 antagonizes NAADP, a second messenger that mediates opening of two pore channels (TPCs) which transport Ca2+ from the endo-lysosomal compartment into the cytosol (48). Hence, a HeLa cell pool depleted of TPC1 (Supp. Figure 1, K) substantially reduced invasion when compared to unmodified HeLa cells. This was particularly pronounced after a 10 min infection pulse, suggesting that Ca2+-dependent invasion is important early in infection (Figure 1, E).

Revised version: Ned19 antagonizes NAADP, a second messenger that mediates opening of two pore channels (TPCs), which transport Ca2+ from the endo-lysosomal compartment into the cytosol (55). Hence, a HeLa cell pool depleted of TPC1 (Supp. Figure 1, G) showed substantially reduced invasion when compared to wildtype HeLa cells. This was particularly pronounced after a 10 min infection pulse, suggesting that Ca2+-dependent invasion is important early in infection (Figure 1, C).

Earlier version: CD38 is a well-known producer of NAADP in immune cells (49–51). Thus, we tested S. aureus invasion in HeLa cells treated with the specific CD38 inhibitor 78c, even though expression of CD38 in these cells is predicted to be very low (proteinatlas.org). We observed no effect on S. aureus invasion by CD38 inhibition (Supp. Figure 1, L). An alternative NAADP producer, with higher expression in our infection model, is Sterile Alpha and TIR Motif Containing 1 [SARM1, (52)]. We measured a slightly decreased invasion efficiency 10 min p.i. in a cell pool lacking SARM1 when compared to WT cells (Figure 1, F; Supp. Figure 1, M), suggesting that SARM1 might be enrolled in NAADP production during internalization of S. aureus by host cells.

Revised version: CD38 is a well-known producer of NAADP in immune cells (56-58). Thus, we tested S. aureus invasion in HeLa cells treated with the specific CD38 inhibitor 78c, even though expression of CD38 in these cells is predicted to be very low (proteinatlas.org). We observed no effect on S. aureus invasion by CD38 inhibition (Supp. Figure 1, H). An alternative NAADP producer, with higher expression in our infection model, is Sterile Alpha and TIR Motif Containing 1 [SARM1, (59)]. We measured a slightly decreased invasion efficiency 10 min p.i. in a cell pool lacking SARM1 when compared to wildtype cells (Figure 1, D; Supp. Figure 1, I), suggesting that SARM1 might be involved in NAADP production during internalization of S. aureus by host cells. However, there might be other enzymes that produce NAADP (60) during S. aureus infection.

Earlier version: Altogether, our findings support the important role of lysosomal Ca2+ mobilization in S. aureus invasion as particularly required early in infection.

Revised version: Altogether, our findings show that interference with lysosomal Ca2+ mobilization limits S. aureus invasion particularly early in infection.

Earlier version: Rapid S. aureus invasion requires lysosomal exocytosis, ASM and its substrate sphingomyelin on the host cell surface

Revised version: S. aureus invasion requires lysosomal exocytosis, ASM and its substrate sphingomyelin on the host cell surface

Added in the revised version: Since our data suggest that lysosomal Ca2+ release facilitates the internalization of S. aureus by host cells, we speculated that lysosomal exocytosis is involved in bacterial invasion. Lysosomal exocytosis is a Ca2+-dependent process during which lysosomes fuse with the plasma membranes and release their content. The internalization of several human pathogens has been shown to depend on lysosomal exocytosis (8, 16, 61-63).

Added in the revised version: To visualize lysosomal exocytosis, we developed an assay that makes use of the split NanoLuc luciferase system [Figure 2, A; (64, 65)]. We here engineered HeLa cells to express the high-affinity NanoLuc peptide HiBiT between the signal peptide and the mature chain of lysosomal-associated membrane protein 1 (LAMP1). Transient localization of LAMP1 to the cell surface of mammalian cells is a hall mark for lysosomal exocytosis (66). To measure the release of the HiBiT-tagged LAMP1, we added LgBit together with the NanoLuc substrate Furimazine to the cell culture medium and measured luminescence of reconstituted NanoLuc in a microplate reader. Maximum luminescence was detected ∼10-15 min upon starting the assay and decline afterwards, likely due to exhaustion of the luciferase substrate (Figure 2, B). Treatment with the Ca2+ ionophore ionomycin, a known activator of lysosomal exocytosis (11, 67), increased the luminescence, whereas Vacuolin-1, an inhibitor of lysosomal exocytosis (68), decreased luminescence when compared to an untreated control. Thus, we concluded that our assay successfully quantified exposure of LAMP1 on the host cell surface.

Removed in the revised version: Ionomycin treatment results in lysosomal exocytosis, a Ca2+-dependent process in which lysosomes release their content to the extracellular space (45, 53). Hence, we treated host cells with Vacuolin-1, an inhibitor blocking this process (54), which reduced invasion efficiency in HuLEC (Figure 2, A) and HeLa cells (Supp. Figure 2, A), while not affecting survival of the bacteria (Supp. Figure 2, B).

S. aureus invasion requires lysosomal exocytosis, ASM and plasma membrane sphingomyelin.

Earlier version: (A, B) S. aureus invasion requires Syt7-dependent lysosomal exocytosis. HuLEC were treated with the lysosomal exocytosis inhibitor Vacuolin-1 and invasion efficiency of S. aureus was determined 30 min p.i. (A, n=3). Invasion efficiency of S. aureus JE2 was determined in HeLa Syt7 K.O. cells 10 min and 30 min p.i. (B, n=5; data normalized to HeLa WT). (C, D) S. aureus internalization is dependent on ASM activity in endothelial and epithelial cells. Indicated cell lines were treated with the ASM inhibitors amitriptyline or ARC39 (C) and invasion efficiency of S. aureus was determined (n≥3). (D) HuLEC were treated with PCK310 or ARC39 for the indicated periods. Subsequently, cells were infected with S. aureus for 30 min and invasion efficiency was determined by CFU counting (n≥4). (E, F) S. aureus invasion requires SM and SMase activity on host cell plasma membranes. HuLEC were treated with the bacterial SMase β-toxin for 75 min and were subsequently infected with S. aureus JE2 for 30 min (n≥3). (F) Presence of extracellular SMase activity restores invasion defect in TPC1 and Syt7 K.O. cell lines. HeLa wild type as well as TPC1 or Syt7 KO cell lines were infected with S. aureus JE2 in presence of 100 ng/ml of the bacterial SMase β-toxin and invasion efficiency 10 min p.i. was determined (n=4). In each experiment the numbers of intracellular S. aureus JE2 were determined by lysostaphin protection assay and CFU counting. Results obtained for the tested condition were normalized to the wildtype cell line or the mock-treated control. (G) scheme of lysosomal exocytosis and ASM release as well as interfering agents. Cer: ceramide, SM: sphingomyelin, ASM: acid sphingomyelinase. Statistics: one sample t-test (A-C, E), one-way ANOVA and Dunnett’s multiple comparisons test (D). Two-way ANOVA and Tukey’s multiple comparison (F). Bars represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001.

Revised version: (A, B) Luminescence-based lysosomal exocytosis assay. HeLa cells expressing LAMP1-HiBit were treated with 1 µM ionomycin (immediately before the measurement) or 1 µM Vacuolin-1 (75 min before the measurement) in presence of LgBit and NanoLuc substrate. Lysosomal exocytosis was monitored by measuring chemiluminescence in a Tecan microplate reader. n=3 (C) S. aureus triggers lysosomal exocytosis during infection. HeLa cells expressing LAMP1-HiBit were pretreated for 75 min with 1 µM Vacuolin-1 or 200 µM trans-Ned19. Subsequently, cells were infected with S. aureus JE2 (MOI10), treated with 1 µM ionomycin or left untreated and lysosomal exocytosis was determined by measuring chemiluminescence in a Tecan microplate reader 10 min p.i. Luminescence detected in infected samples was scaled to untreated controls (0 %) and ionomycin (100%) to determine the proportion of all “releasable” lysosomes liberated during infection. n=6. (D-F) S. aureus invasion requires Syt7-dependent lysosomal exocytosis. HuLEC were treated with the lysosomal exocytosis inhibitor Vacuolin-1 or the lysosomal exocytosis inducer ionomycin and invasion efficiency of S. aureus was determined 30 min p.i. (D, E, n≥3). Invasion efficiency of S. aureus JE2 was determined in HeLa Syt7 K.O. cells 10 min and 30 min p.i. (F, n≥5; data normalized to HeLa wildtype). (G, H) S. aureus internalization is dependent on ASM activity in endothelial and epithelial cells. Indicated cell lines were treated with the ASM inhibitors amitriptyline or ARC39 (G) and invasion efficiency of S. aureus was determined (n≥3). (G) HuLEC were treated with PCK310 or ARC39 for the indicated periods. Subsequently, cells were infected with S. aureus for 30 min and invasion efficiency was determined by CFU counting (H). n≥4. (I, J) Genetic ablation of ASM renders S. aureus invasion insensitive towards ARC39. HeLa ASM K.O. cells were treated with 10 µM ARC39 or were left untreated and the number of invaded bacteria was determined 30 min p.i.. Invasion efficiency was determined by normalization to wildtype HeLa cells (I) or to corresponding untreated samples (J, n≥4), (K) Genetic ablation of ASM is accompanied by massive alteration in cellular sphingolipid profiles. HeLa wildtype and ASM K.O. cells were treated with 20 µM amitriptyline (75 min) and 10 µM ARC39 (22h) or left untreated, as indicated. Whole cell sphingolipid profiles were determined by HPLC-MS/MS and the ratios of SM vs. Cer were calculated for lipid species with varying acyl side chains (chain length indicated on the x-axis). n= 4 (L) S. aureus invasion requires SM on host cell plasma membranes. HuLEC were treated with the bacterial SMase β-toxin for 75 min and were subsequently infected with S. aureus JE2 for 30 min (n≥3). (M) Experimental design for β-toxin treatment of host cells during S. aureus infection. Host cells were either pretreated with β-toxin to remove SM from the plasma membrane prior to infection (upper panel; Figure 2, L). Alternatively, β-toxin was added together with the bacteria to rescue the absence of ASM (lower panel, Figure 2, N). (N) Presence of extracellular SMase activity restores the invasion defect in TPC1 and Syt7 K.O. cell lines. HeLa wild type as well as TPC1 or Syt7 KO cell lines were infected with S. aureus JE2 in presence of 100 ng/ml of the bacterial SMase β-toxin and invasion efficiency 10 min p.i. was determined (n=4). In each experiment the numbers of intracellular S. aureus JE2 were determined by lysostaphin protection assay and CFU counting. Results obtained for the tested condition were normalized to the wildtype cell line or the mock-treated control. (O) Scheme of lysosomal exocytosis and ASM release as well as interfering agents. Cer: ceramide, SM: sphingomyelin, ASM: acid sphingomyelinase. xStatistics: one sample t-test (D-G, L), one-way ANOVA and Dunnett’s multiple comparisons test (C, J). Mixed-effects model (REML) and Šídák’s multiple comparisons (H). Two-way ANOVA and Tukey’s multiple comparison (K). Two-way RM ANOVA and Šídák’s multiple comparison (N) Bars represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001. Created in BioRender.

Added in the revised version: Next, we monitored LAMP1 surface levels during S. aureus infection. We assumed that ionomycin would trigger exocytosis of all “releasable” lysosomes. Hence, the detected luminescence signal during infection was scaled to an untreated control (=0 %) as well as ionomycin (=100%), to express LAMP1 surface levels as percentage of all “releasable” lysosomes. We detected that ∼30% of all “releasable” lysosomes were liberated from the reporter cells 10 min after infection with S. aureus (Figure 2, C). Pretreatment of host cells with Vacuolin-1 or Ned-19 strongly reduced LAMP1 surface levels, suggesting that S. aureus triggers lysosomal exocytosis in an NAADP-dependent manner during invasion.

Added in the revised version: To test whether lysosomal exocytosis is required for S. aureus internalization by host cells, we measured invasion efficiency in HuLEC (Figure 2, D) and HeLa cells (Supp. Figure 2, A) pretreated with Vacuolin-1. Invasion was strongly reduced in inhibitor-treated samples, while the survival of the bacteria was not affected by the compound (Supp. Figure 2, B).

Added in the revised version: Next, we triggered lysosomal exocytosis in HuLEC by addition of ionomycin prior to infection, thereby depleting the pool of “releasable” lysosomes and thus preventing lysosomal exocytosis during infection. Ionomycin pretreatment reduced invasion efficiency in a concentration-dependent manner (Figure 2, E). Bacterial growth and survival as well as host cell integrity were not affected at the concentrations used (1 µM, Supp. Figure 2, C-E).

Removed in the revised version: Created with BioRender.com/u99k818.

Earlier version: Synaptotagmin 7 (Syt7) is known to support fusion of lysosomes with the plasma membrane in a Ca2+-dependent manner (55). Accordingly, Syt7-depleted HeLa cells demonstrated lower bacterial invasion (Figure 2, B). As already observed for TPC1 (Figure 1, E), this was most pronounced early in infection further supporting our hypothesis that lysosomal exocytosis is important for rapid invasion events.

Revised version: Synaptotagmin 7 (Syt7) is known to support fusion of lysosomes with the plasma membrane in a Ca2+-dependent manner (69). Accordingly, Syt7-depleted HeLa cells demonstrated lower bacterial invasion (Figure 2, F). As already observed for TPC1 (Figure 1, C), this was most pronounced early in infection. Taken together, our results support the hypothesis that lysosomal exocytosis is important for invasion of host cells by S. aureus particularly early during infection.

Earlier version: Lysosomal exocytosis results in release of ASM (45, 56). Therefore, prior to infection, the cells were exposed to amitriptyline, a functional inhibitor of ASM (FIASMA), or a competitive inhibitor [ARC39, (57)] at concentrations not affecting bacterial survival (Supp. Figure 2, C, D). While invasion in HuLEC, human umbilical vein endothelial cells (HuVEC) and HeLa was sensitive to both inhibitors, invasion in 16HBE14o- and EA.hy926 was not or only marginally affected (Figure 2, C). We excluded that differences among cell lines arose from differential ASM activity or sensitivity to the inhibitors by monitoring enzymatic activity [Supp. Figure 2, E; (58) and Supp. Figure 2, F; (59)].

Revised version: Lysosomal exocytosis results in release of ASM (6, 11), an enzyme which has previously been associated with the internalization of several bacterial and viral pathogens (8, 13, 16, 62, 63, 70) and thus, we tested if ASM activity is also required for the uptake of S. aureus by host cells. Therefore, we treated different cell lines prior to infection with amitriptyline, a functional inhibitor of ASM (FIASMA), or the competitive ASM inhibitor ARC39 (71) and infected host cells with S. aureus. While invasion in HuLEC, human umbilical vein endothelial cells (HuVEC) and HeLa was reduced by both inhibitors, invasion in bronchial epithelial 16HBE14o- and the endothelial-epithelial hybrid cell line EA.hy926 was not or only marginally altered (Figure 2, G). The used inhibitor concentrations did not affect bacterial survival (Supp. Figure 2, F, G). Furthermore, we excluded that differences among cell lines arose from differential ASM activity or sensitivity to the inhibitors by monitoring enzymatic activity [Supp. Figure 2, H; (72) and Supp. Figure 2, I; (73)].

Earlier version: Next, we treated HuLEC with 10 µM ARC39 or 0.5 µM PCK310, a fast-acting ASM inhibitor. Already after 1h preincubation, PCK310 reduced invasion of S. aureus, whereas ARC39 required overnight treatment for a similar reduction (Figure 2, D). Taken together, these results show that ASM is involved in S. aureus invasion of host cells.

Revised version: Next, we treated HuLEC with 10 µM ARC39 or 0.5 µM PCK310, a fast-acting ASM inhibitor. Already after 1h preincubation, PCK310 reduced invasion of S. aureus, whereas ARC39 required overnight treatment for a similar reduction (Figure 2, H). We confirmed the reduction of internalized colony forming units (CFU) by amitriptyline and PCK310 treatment with a microscopy-based approach, where we determined the number of invaded fluorescent bacteria per host cell (Supp. Figure 2, J). Taking together, our results show that ASM is involved in S. aureus invasion of host cells in a cell-type specific manner.

Added in the revised version: Next, we generated two ASM-depleted HeLa cell pools using CRISPR/Cas9 (74) with two distinct small guide RNAs (sgRNAs) and determined successful ablation of cellular ASM activity with a visible range ASM FRET probe [(73), Supp. Figure 2, K]. Surprisingly, we did not detect a reduced invasion efficiency in ASM K.O. cells, when compared to wildtype (Figure 2, I). To test if unspecific side effects of the ASM inhibitors could have contributed to reduced invasion in our previous experiments, we treated wildtype and ASM K.O.s with ARC39 and tested for S. aureus invasion efficiency (Figure 2, K). In wildtype cells, bacterial invasion again was strongly reduced after ARC39 treatment, whereas only a slight reduction was observed in the K.O.s, which may represent unspecific side effects of the inhibitor. Since the reducing effect of ARC39 on invasion is almost completely lost in ASM K.O. cells, the reduced invasion in wildtype cells upon ARC39 treatment predominantly must be caused by the inhibition of ASM.

Added in the revised version: Absence of ASM results in cellular accumulation of SM and usually is accompanied by the neurodegenerative diseases, Niemann-Pick syndrome A and B (75). Hence, we compared the sphingolipidomes of inhibitor-treated wildtype and ASM K.O. cells by high-pressure liquid chromatography coupled to tandem mass spectrometry (HPLC-MS/MS). The SM accumulation is described as ratio of cellular SM vs. ceramide concentrations (ASM educt vs. product) for lipid species with different acyl chain lengths (Figure 2, K). Whereas we did observe a moderately increased amount of most SM species upon ARC39 treatment, cells exposed to amitriptyline only showed a slight increase in SM 16:0, which can be explained by the shorter preincubation time used for amitriptyline (75 min vs. 22h). However, in both ASM K.O. cell lines we observed markedly increased SM levels when compared to either untreated or inhibitor-treated wildtype cells. The relatively low increase of SM levels after inhibitor treatment likely arises from short inhibitor pulses (in the range of minutes or hours), whereas SM in ASM K.O. cell lines accumulates over the entire culture period (days or weeks since gene ablation). We hypothesize that in the latter situation the host cell invasion of S. aureus is generally increased, while the ASM-dependent uptake pathway is absent. This is supported by the high residual internalization of S. aureus by ASM K.O. cells upon ARC39 treatment. Previously, it was shown that S. aureus host cell entry is limited by caveolin-1, a protein associated with lipid microdomains (76, 77). Absence of ASM was demonstrated to interfere with caveolin-1-associateded endocytosis (78, 79) and thus, the increased invasion in ASM K.O. cells may be caused by dysfunctional caveolin-1.

Earlier version: We removed SM, the substrate of ASM, from plasma membranes of HuLEC (Figure 2, E) and HeLa (Supp. Figure 2, G) by pretreatment with the bacterial SMase β-toxin. Most human pathogenic S. aureus strains including JE2, the strain we use in the present study, do not produce β-toxin, due to a genomic integration of a prophage (60). SMase pretreatment reduced invasion efficiency in a concentration-dependent manner in both cell lines. 10 ng/ml β-toxin were sufficient to decrease invasion to 34.37 ± 6.24% of that measured in untreated control cells. Interestingly, a 250-fold higher concentration did not lead to further reduction. In line with our previous experiments, this suggests that multiple S. aureus invasion pathways exist within the same host cell population, of which at least one requires SM and its enzymatic breakdown on the plasma membrane.

Revised version: Release of ASM by host cells results in cleavage of SM on the host cell surface, a process that previously was suggested to facilitate cell entry of human pathogens (7, 8, 13-17). To test whether S. aureus invasion requires SMase activity on the host cell surface, we removed SM, the substrate of ASM, from plasma membranes of HuLEC (Figure 2, L) and HeLa (Supp. Figure 2, L) by pretreatment with the bacterial SMase β-toxin (see Figure 2, M). We confirmed β-toxin-dependent SM conversion in living cells by detecting the ceramide metabolites of fluorescent SM analogs [Supp. Figure 2, M; (81)]. Plasma membrane pretreatment with bacterial SMase reduced S. aureus invasion efficiency in a concentration-dependent manner in both cell lines. 10 ng/ml β-toxin were sufficient to decrease invasion to ∼30% of that measured in untreated control cells. Interestingly, a 250-fold higher concentration did not lead to further reduction, suggesting that 10 ng/ml β-toxin were sufficient to quantitatively convert SM to ceramide at the cell surface. In line with our previous experiments, this suggests that multiple S. aureus invasion pathways exist within the same host cell population, of which at least one requires SM and its enzymatic breakdown on the plasma membrane.

Earlier version: Next, we infected wild-type as well as the TPC1 or Syt7 K.O. cells with S. aureus in presence (but without pretreatment) of 100 ng/ml β-toxin and determined the number of intracellular bacteria after 10 min (Figure 2, F). Presence of β-toxin in the culture medium completely rescued the invasion defect of the K.O. cells. This indicates that reduced bacterial internalization in absence of TPC1 and Syt7 results from a lack in ASM delivery to the cell surface by Ca2+-dependent lysosomal exocytosis.

Revised version: Since we detected enhanced Ned19-sensitive lysosomal exocytosis during S. aureus infection (Figure 2, C), we hypothesized that absence of TPC1 or Syt7 should impair lysosomal exocytosis and thereby release of ASM. Therefore, we infected wildtype as well as TPC1 or Syt7 K.O. cells with S. aureus in presence (but without pretreatment, see Figure 2, M) of 100 ng/ml β-toxin and determined the number of intracellular bacteria after 10 min (Figure 2, N). The addition of β-toxin to the culture medium completely rescued the invasion defect of the K.O. cells. This indicates that reduced bacterial internalization in absence of TPC1 and Syt7 results from a lack in ASM delivery to the cell surface by Ca2+-dependent lysosomal exocytosis and that absence of ASM from the host cell surface, can be compensated by exogenous addition of the bacterial SMase β-toxin.

Added in the revised version: We want to emphasize the difference in experimental application of β-toxin (Figure 2, M) where we i) pretreat host cells with the bSMase to remove SM from the plasma membrane (Figure 2, L and Supp. Figure 2, L) and ii) rescue the absence of exocytosed ASM from the host cell surface by exogenous addition of β-toxin during the infection (Figure 2, N). In the former case, pretreatment leads to removal of SM from the plasma membrane thus preventing SM conversion by ASM during infection and resulting in reduced invasion. The latter leads to SM conversion by the bacterial SMase during infection thereby increasing invasion.

Added in the revised version: In summary, our data suggest the existence of a S. aureus invasion pathway that depends on the metabolic breakdown of SM by ASM on the host cell surface (Figure 2, O).

ASM- and Ca2+-mediated invasion is rapid

Earlier version: To study the role of SM in S. aureus uptake kinetics, we incubated HuLEC with fluorescent sphingolipid analogs and recorded S. aureus infection by time lapse imaging [Figure 3, A; Supp. Video 1; Supp. Figure 3; Supp. Video 2; (59)]. We observed that the bacteria were rapidly engulfed by SM-containing membrane compartments. While association of bacteria with BODIPY-FL-C12-SM increased over time during the first 30 min of infection, this was not observed for BODIPY-FL-C12-Ceramide, suggesting that S. aureus invasion specifically relies on interaction with SM in plasma membranes (Figure 3, B). Next, we determined the invasion efficiency of S. aureus in a time-dependent manner and blocked invasion by either ASM inhibitors, ionomycin, BAPTA-AM or β-toxin pretreatment. We found a time-dependent increase in intracellular bacteria, whereby all treatment conditions reduced bacterial invasion at every time point with the untreated control at 30 min set to 100% (Figure 3, C). By normalizing the data to the untreated control at the corresponding time points, we found that -with exception of β-toxin-the effect of the treatments was most pronounced early (5-10 min) after infection (Figure 3, D). Thus, early during infection bacteria enter the host cell predominantly in an ASM- and Ca2+-dependent manner.

Revised version: To study the role of SM in S. aureus uptake kinetics, we incubated HuLEC with fluorescent sphingolipid analogs and recorded S. aureus infection by time lapse imaging [Figure 3, A; Supp. Video 1; Supp. Figure 3; Supp. Video 2; (73)]. We observed that the bacteria were rapidly engulfed by SM-containing membrane compartments. While association of bacteria with BODIPY-FL-C12-SM increased over time during the first 30 min of infection, this was not observed for BODIPY-FL-C12-Ceramide, suggesting that S. aureus invasion specifically relies on interaction with SM in host cell plasma membranes (Figure 3, B).

ASM- and Ca2+-dependent uptake is rapid and predominantly mediates invasion early in infection.

Earlier version: (A, B) S. aureus associates with SM early during invasion. HuLEC were pretreated with BODIPY-FL-C12-SM or BODIPY-FL-C12-ceramide and infected with red-fluorescent S. aureus. Scale bar: 5 µm. Infection was monitored by live cell imaging and the proportion of bacteria that associated with the lipid analogs was quantified in each individual frame (SM: n=5, Cer: n=3). (C, D) ASM and Ca2+-dependent invasion is rapid. HuLEC were treated with ARC39, amitriptyline, BAPTA-AM, ionomycin, or the bacterial SMase β-toxin. Then, cells were infected with S. aureus and the number of invaded bacteria was determined after different time periods. The results were either normalized to 30 min time point of untreated controls (C) or to the corresponding time points of the untreated controls (D). (n ≥5). Statistics: Mixed effect analysis and Tukey‘s multiple comparison. Graphs represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001.

Revised version: (A, B) S. aureus associates with SM early during invasion. HuLEC were pretreated with BODIPY-FL-C12-SM or BODIPY-FL-C12-ceramide and infected with red-fluorescent S. aureus. Scale bar: 5 µm. Infection was monitored by live cell imaging and the proportion of bacteria that associated with the lipid analogs was quantified in each individual frame (SM: n=5, Cer: n=3). (C, D) ASM and Ca2+-dependent invasion is rapid. HuLEC were treated with ARC39, amitriptyline, BAPTA-AM, ionomycin, or the bacterial SMase β-toxin. Then, cells were infected with S. aureus and the number of invaded bacteria was determined after different time periods. The results were either normalized to the 30 min time point of untreated controls (C) or to the corresponding time points of the untreated controls (D). (n ≥5). Statistics: Mixed effect analysis and Tukey‘s multiple comparison. Graphs represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001. Created in BioRender.

Added in the revised version: Next, we determined the invasion efficiency of S. aureus in a time-dependent manner and blocked invasion by either ASM inhibitors, ionomycin, BAPTA-AM or β-toxin pretreatment. We found a time-dependent increase in intracellular bacteria, whereby all treatment conditions reduced bacterial invasion at every time point with the untreated control at 30 min set to 100% (Figure 3, C). We next normalized the data to the untreated control of the corresponding time point. Except for the β-toxin treatment, the effect of all other treatments was most pronounced early (5-10 min) after infection (Figure 3, D). For instance, ionomycin treatment massively reduced S. aureus invasion to ∼20% of untreated controls when we infected for 10 min. We concluded that in untreated controls 80% of the bacteria entered host cells via the pathway that depends on lysosomal exocytosis, Ca2+ and ASM, which is blocked by ionomycin treatment. Consequently, bacteria that were able to invade host cells in ionomycin-treated samples (the residual ∼20%) must have employed other pathways. By contrast, in samples infected for 30 min a higher proportion of bacteria was able to enter host cells despite ionomycin treatment (∼50% compared to untreated controls). This suggests that only 50% of all invaded bacteria were internalized in an ASM- and Ca2+-dependent fashion upon the longer infection pulse. Together with the observed invasion defects in host cells with gene deletions in TPC1 (Figure 1, C) and Syt7 (Figure 2, F), which we exclusively detected during short infection times, our data suggest that early during infection bacteria enter host cells predominantly in an ASM- and Ca2+-dependent manner. Thus, we concluded that the internalization pathway that depends on lysosomal exocytosis, ASM and Ca2+ is rapid when compared to other host cell entry pathways.

Removed in the revised version: Created with BioRender.com/f02t628.

The mode of S. aureus invasion affects infection outcome

Earlier version: After invasion, S. aureus resides within Rab5-positive early and subsequently, in Rab7-positive late phagosomes (33). We pretreated a HeLa cell line stably expressing mCherry-Rab5 as well as YFP-Rab7 with amitriptyline and infected with S. aureus for different time periods and determined the proportion of bacteria that associated with Rab5 and/or Rab7 (Figure 4, A).

Revised version: Next, we tested whether the pathway used for host cell entry also affects the intracellular fate of S. aureus. After invasion, the bacteria reside within Rab5-positive early and subsequently, in Rab7-positive late phagoendosomes (39). To trace phagosome maturation, we generated a HeLa cell line stably expressing mCherry-Rab5 as well as YFP-Rab7 and infected with S. aureus. We observed that the bacteria transiently associate with Rab5 for a few minutes, before they become positive for Rab7 (Supp. Video 3).

Removed in the revised version: The intracellular fate of S. aureus is determined by the pathway of cell entry.

Removed in the revised version: (A) Inhibition of ASM delays formation of Rab7-positive phagosomes. HeLa cells expressing mCherry-Rab5 and YFP-Rab7 were either treated with amitriptyline or left untreated. Then, cells were infected with S. aureus JE2 for indicated periods. Extracellular bacteria were removed and proportion of intracellular bacteria associating with Rab5 or Rab7 was determined by CLSM. Results obtained for amitriptyline-treated samples were normalized to untreated controls (n=5). (B) Detection of phagosomal SM and phagosomal escape by a reporter cell line expressing LyseninW20A -YFP and RFP-CWT. After internalization, S. aureus resides in a phagosome preventing the recruitment of RFP-CWT to S. aureus cell wall and LyseninW20A-YFP to luminal SM, respectively. When the bacteria lyse the phagosomal membrane, luminal SM gets exposed to the cytosol and attracts LyseninW20A -YFP, while RFP-CWT is recruited to the staphylococcal surface. (C) Blocking ASM-dependent internalization affects phagosomal escape. RFP-CWT and LyseninW20A -YFP expressing HeLa were treated with PCK310, or ARC39 and infected with S. aureus strains JE2 or Cowan I. Thereby, β-toxin was either removed prior to infection (pretreatment only) or was present during the whole experiment (pretreatment + during infection). By CLSM, proportions of bacteria that recruited RFP-CWT (phagosomal escape) were determined (n≥3). (D-F) Phagosomal escape depends on presence of plasma membrane SM during invasion, but not presence of SM within phagosomal membranes HeLa RFP-CWT LyseninW20A -YFP were pretreated with β-toxin to remove surface SM (2a, 2b) or left untreated (2c). Then, cells were infected with S. aureus JE2 in presence (3a) or absence (3b) of β-toxin. Untreated samples (3c) were infected with S. aureus JE2 harboring a plasmid either encoding β-toxin and the fluorescence protein Cerulean (pCer+hlb) or solely Cerulean (pCer). Proportion of bacteria that recruited RFP-CWT (phagosomal escape, E) and the percentage of phagosomal escape events that additionally were positive for LyseninW20A-YFP (F) were determined at indicated time points p.i. (n=5). (G, H) Early ASM-dependent invaders possess lower escape rates than late invaders. HeLa cells expressing RFP-CWT were infected with indicated MOIs of S. aureus JE2 either for 10 min (early invaders) or 30 min (early+late invaders). Phagosomal escape rates were determined 3h p.i. (G). HeLa reporter cells expressing YFP-CWT were infected with an MOI=5 of S. aureus JE2 expressing a fluorescent protein (e.g. RFP) for 30 min (early+late invaders). After 20 min, the same samples were infected with S. aureus JE2 expressing another fluorophore (e.g. Cerulean) for 10 min (early invaders) and phagosomal escape was determined 3 h p.i. (H). Statistics: one sample t-test (A), Mixed effects analysis (REML) and Tukey’s multiple comparison (C,G, H), Two-way ANOVA and Šídák’s multiple comparisons test (E, F). Graphs represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001.

Removed in the revised version: Created with BioRender.com/f97r445.

Earlier version: While ASM inhibition marginally affected S. aureus association with Rab5-positive phagosomes, that with the Rab7-positive compartment was significantly reduced compared to the untreated control. This was restricted to shorter infection pulses (5-45 min) and was not observed for long infection periods (60 min). We also observed a reduced proportion of bacteria associating with Rab7-positive membranes after treatment with PCK310 or bacterial SMase. However, this was not caused by translocation of S. aureus to the host cytosol (Supp. Figure 4, A, B). Hence, phagosomes generated ASM-dependently possess different maturation dynamics when compared to phagosomes formed in an ASM-independent fashion.

Revised version: We next blocked the ASM-dependent invasion pathway by amitriptyline treatment and infected with S. aureus for different time periods. We determined the proportion of bacteria that was associated with Rab5 and/or Rab7 (Figure 4, A; see Supp. Figure 4 for bacteria proportion that associates with Rab7). While ASM inhibition only marginally affected S. aureus association with Rab5-positive phagosomes, the proportion of bacteria residing in Rab7-positive vesicles was significantly reduced when compared to the untreated control. This observation was restricted to shorter infection pulses (5-45 min) and was not observed for long infection periods (60 min).

Added in the revised version: Blocking ASM-dependent invasion affects phagosomal maturation and escape during S. aureus infection

Added in the revised version: (A) Inhibition of ASM delays formation of Rab7-positive phagosomes. HeLa cells expressing mCherry-Rab5 and YFP-Rab7 were either treated with amitriptyline or left untreated. Then, cells were infected with S. aureus JE2 for indicated periods. Extracellular bacteria were removed and percentage of intracellular bacteria which associated with Rab5 or Rab7 was determined by CLSM. Results obtained for amitriptyline-treated samples were normalized to untreated controls (n=5). (B) Detection of phagosomal SM and phagosomal escape by a reporter cell line expressing LyseninW20A -YFP and RFP-CWT. After internalization, S. aureus resides in a phagosome preventing the recruitment of RFP-CWT to S. aureus cell wall and LyseninW20A-YFP to luminal SM, respectively. When the bacteria lyse the phagosomal membrane, luminal SM gets exposed to the cytosol and attracts LyseninW20A -YFP, while RFP-CWT is recruited to the staphylococcal surface. (C-F) Blocking ASM-dependent internalization affects phagosomal escape. RFP-CWT and LyseninW20A -YFP expressing HeLa were treated with PCK310,ARC39 (C, F), amitriptyline (D) or Vacuolin-1 (E) and infected with S. aureus strains JE2 or Cowan I. By CLSM, proportions of bacteria that recruited RFP-CWT (phagosomal escape) were determined 3h p.i. if not indicated otherwise (n≥3). Statistics: one sample t-test (A), Mixed effects analysis (REML) and Tukey’s multiple comparison (C,), paired Student’s t-test (D, E) Graphs represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001. Created in BioRender.

Added in the revised version: We also observed a reduced proportion of bacteria associated with Rab7-positive membranes after treatment with PCK310 or bacterial SMase (Supp. Figure 4, B, C). Forty-five minutes p.i., about half of the bacteria were localized to Rab7-positive vesicles in untreated cells, whereas only ∼40% associated with Rab7 upon amitriptyline, PCK310 or β-toxin treatment (Supp. Figure 4, C). However, this was not caused by translocation of S. aureus to the host cytosol, which was excluded by expressing the fluorescence reporter RFP-CWT (42) in the YFP-Rab7 cell line. Accordingly, RFP-CWT-labelled bacteria in the host cytosol were excluded from the data set [Supp. Figure 4, B-D, (42)]. Hence, ASM-dependently generated phagoendosomes possess different maturation dynamics when compared to phagosomes formed in an ASM-independent fashion.

Added in the revised version: Next, we measured phagosomal escape while simultaneously detecting SM content of disrupted phagosomal membranes. Therefore, we used a HeLa reporter cell line that cytosolically expresses the phagosomal escape reporter RFP-CWT (42) as well as the SM reporter LyseninW20A-YFP (82), respectively (Figure 4, B). To validate the reporter system, we removed SM from the plasma membrane by β-toxin treatment prior to infection and monitored phagosomal escape by live cell imaging (Supp. Figure 5, Supp. Video 4). Bacterial translocation into the host cytosol is, again, indicated by recruitment of RFP-CWT. In untreated cells, phagosomal escape was accompanied by recruitment of LyseninW20A-YFP to vesicular membranes, suggesting that SM located at the luminal leaflet of phagosomes was exposed to the cytosol. By contrast, LyseninW20A-YFP was not recruited to escape events in β-toxin-treated cells, indicating that the SMase had removed SM from the membranes.

Earlier version: Next, we measured phagosomal escape while simultaneously detecting SM content of disrupted phagosomal membranes. Therefore, we used a HeLa reporter cell line that cytosolically expresses the phagosomal escape reporter RFP-CWT (36) as well as the SM reporter LyseninW20A-YFP (61), respectively (Figure 4, B). Reporter gene expression in the cell line did not alter S. aureus invasion dynamics under different treatment conditions (Supp. Figure 4, C). The reporter cells were left untreated or were pre-exposed to ASM inhibitors and were infected with S. aureus strain JE2. In addition, S. aureus Cowan I, a strain known to have low escape rates, was used for infection. The proportion of bacteria that escaped from the phagosome (Figure 4, C) as well as the proportion of escape events that additionally were positive for cytosolic SM (Supp. Figure 4, D) were determined 1.5 h and 3 h p.i. (see also Supp. Figure 5, A-D).

Revised version: Next, we tested whether internalization of the bacteria via the ASM-dependent pathway would affect phagosomal escape of S. aureus. We first confirmed that invasion of S. aureus is sensitive to ASM inhibitors and β-toxin treatment, despite reporter gene expression in the cell line (Supp. Figure 4, E). Then reporter cells either were left untreated or were pre-exposed to ASM inhibitors and were infected with S. aureus strain JE2. In addition, S. aureus Cowan I, a non-cytotoxic strain known to have low escape rates, was used for infection. The proportion of bacteria that escaped from the phagosome (Figure 4, C, D) as well as the proportion of escape events, which were additionally positive for cytosolic SM (Supp. Figure 4, F) were determined 1.5 h and 3 h p.i. (see also Supp. Figure 6, A-D). Whereas no differences for Lysenin-positive escape events were detected (Supp. Figure 4, F), a higher proportion of bacteria escaped from phagosomes on abrogation of ASM activity regardless of the inhibitor used (Figure 4, C, D). A similar increase was observed when we blocked lysosomal exocytosis by Vacuolin-1 ((Figure 4, E), suggesting that the lack of ASM delivery to the host cell surface during bacterial invasion affects phagosomal escape.

Removed in the revised version: Whereas no differences for Lysenin-positive escape events were detected (Supp. Figure 4, D), a higher proportion of bacteria escaped from phagosomes on abrogation of ASM activity regardless of the inhibitor used (Figure 4, C; Supp. Figure 4, E). A similar increase was observed when we blocked lysosomal exocytosis by Vacuolin-1 (Supp. Figure 4, F).

Earlier version: Since blocking of ASM-dependent uptake reduced the number of intracellular bacteria but increased their phagosomal escape rates, we hypothesized that the presence of SM on the plasma membrane affects downstream events during S. aureus infection. To address this question, we pretreated host cells with the bacterial SMase β-toxin to block SM-dependent invasion (Figure 4, D; 2a, 2b) and either infected the cells in presence (3a) or after removal of the toxin (3b). β-toxin pretreatment, resulted in enhanced phagosomal escape rates of S. aureus JE2, independent of the presence of β-toxin during infection (Figure 4, E; 4a, 4b). This was not observed for infection with the non-cytotoxic Cowan I strain, indicating that removal of SM by β-toxin did not affect general phagosomal membrane stability (Supp. Figure 4, G). Moreover, β-toxin treatment strongly reduced recruitment of LyseninW20A, demonstrating effective removal of SM from membranes by β-toxin (Figure 4, F and Supp. Figure 5, E-J).

Revised version: Since blocking of ASM-dependent uptake reduced invasion efficiency yet increased phagosomal escape rates, we hypothesized that the presence of SM on the plasma membrane affects downstream events during S. aureus infection. To address this, we pretreated host cells with the bacterial SMase β-toxin to block SM-dependent invasion (Figure 5, A; 2a, 2b) and either infected the cells in presence (Figure 5, A; 3a) or after removal of the toxin (3b). β-toxin pretreatment, resulted in enhanced phagosomal escape rates of S. aureus JE2, independent of the presence of β-toxin during infection (Figure 5, B; 4a, 4b). This was not observed for infections with the Cowan I strain, indicating that removal of SM by β-toxin did not affect general phagosomal membrane stability (Supp. Figure 4, G). Moreover, β-toxin treatment strongly reduced recruitment of LyseninW20A, demonstrating effective removal of SM from membranes by β-toxin (Figure 5, C and Supp. Figure 6, E-J).

Added in the revised version: The intracellular fate of S. aureus is determined by host cell entry

Added in the revised version: (A-C) Phagosomal escape depends on presence of plasma membrane SM during invasion, but not presence of SM within phagosomal membranes HeLa RFP-CWT LyseninW20A -YFP were pretreated with β-toxin to remove surface SM (2a, 2b) or left untreated (2c). Then, cells were infected with S. aureus JE2 in presence (3a) or absence (3b) of β-toxin. Untreated samples (3c) were infected with S. aureus JE2 harboring a plasmid either encoding β-toxin and the fluorescence protein Cerulean (pCer+hlb) or solely Cerulean (pCer). Proportion of bacteria that recruited RFP-CWT (phagosomal escape, B) and the percentage of phagosomal escape events that additionally were positive for LyseninW20A-YFP (C) were determined at indicated time points p.i. (n=5). (D, E) Early ASM-dependent invaders possess lower escape rates than late invaders. HeLa cells expressing RFP-CWT were infected with indicated MOIs of S. aureus JE2 either for 10 min (early invaders) or 30 min (early+late invaders). Phagosomal escape rates were determined 3h p.i. (D). HeLa reporter cells expressing YFP-CWT were infected with an MOI=5 of S. aureus JE2 expressing a fluorescent protein (e.g. RFP) for 30 min (early+late invaders). After 20 min, the same samples were infected with S. aureus JE2 expressing another fluorophore (e.g. Cerulean) for 10 min (early invaders) and phagosomal escape was determined 3 h p.i. (E). (F) Summary of experiments analyzing the influence of invasion pathway on phagosomal escape. The measured effects of different conditions on rapid ASM-dependent invasion and accompanied alterations in phagosomal escape rates are summarized. red = increased, blue=decreased. Statistics: Two-way ANOVA and Šídák’s multiple comparisons test (B-E). Graphs represent means ± SD. *p≤0.05, **p≤0.01, ***p≤0.001, ****p≤0.0001. Created in BioRender.

Earlier version: To investigate if removal of SM was important at the cell surface or within the phagosome, we infected the cells with an S. aureus JE2 strain recombinantly overexpressing β-toxin and Cerulean, or a strain solely expressing Cerulean. We did not observe increased phagosomal escape upon overexpression of β-toxin (Figure 5, E; 4c), although SM levels of phagosomal membranes were reduced (Figure 5, F) comparably to samples pretreated with β-toxin (see above). Moreover, we did not detect substantial differences in escape rates of S. aureus 6850, a strain naturally producing β-toxin (Supp. Figure 4, H), and an isogenic β-toxin mutant. By contrast, escape of S. aureus 6850 again was enhanced upon pretreatment with β-toxin (Supp. Figure 4, I) accompanied by absence of SM from phagosomal membranes (Supp. Figure 4, J).

Revised version: To investigate if removal of SM was important at the cell surface or within the phagosome, we infected the cells with transgenic S. aureus JE2 strain either collinearly overexpressing β-toxin and cyan-fluorescent reporter Cerulean or solely expressing Cerulean. We did not observe increased phagosomal escape upon overexpression of β-toxin (Figure 5, B; 4c), although Lysenin recruitment to phagosomal membranes was reduced to levels of β-toxin-pretreated samples (Figure 5, C). Moreover, we did not detect substantial differences in escape rates of S. aureus 6850, a strain naturally producing β-toxin (Supp. Figure 4, H), and its isogenic β-toxin mutant. Again, escape of S. aureus 6850 was enhanced upon pretreatment with β-toxin (Supp. Figure 4, I) accompanied by absence of SM from phagosomal membranes (Supp. Figure 4, J).

Added in the revised version: Phagosomal escape of S. aureus thus is independent of SM levels within the infected phagosomes but rather reflects the presence of SM and ASM at the plasma membrane during host cell entry.

Added in the revised version: We previously showed that S. aureus predominantly enters host cells via the rapid ASM- and Ca2+-dependent pathway early upon host cell contact (Figure 1, C; Figure 2, F; Figure 3, D), while other pathways are likely employed later during infection. Consequently, the proportion of pathways (ASM-dependent vs. ASM-independent) that is used for host cell entry relies on the infection time. To demonstrate that phagosomal escape is dependent on the pathway of cell entry, we infected reporter cells with S. aureus for either a 10 min (“early invaders”) or a 30 min infection pulse (representing “early and late invaders”) and at 3h p.i. we determined phagosomal escape rates (Figure 5, D) as well as the number of invaded bacteria (Supp. Figure 4, K). Independent of the multiplicity of infection (MOI) used, bacteria taken up within the first 10 min demonstrated significantly lower escape rates when compared to bacteria that were cocultured with host cells for 30 min.

Added in the revised version: This was corroborated by HeLa YFP-CWT infected with S. aureus expressing a fluorescent protein (e.g. mRFP) for 30 min (“early and late invaders”). After 20 min, we initiated a second infection pulse for 10 min (“early invaders”) with a S. aureus strain expressing another fluorescence protein (e.g. Cerulean) and determined phagosomal escape rates of both bacterial recombinants (Figure 5, E). Again, early invaders demonstrated lower escape rates when compared to the 30 min control. This was independent of the fluorescent marker expressed by the strain that was used for each infection pulse.

Added in the revised version: Taken together, we observed higher escape rates 3h p.i., whenever we blocked the rapid ASM-dependent invasion pathways. Consistently, escape rates of bacteria predominantly internalized via the ASM-dependent route (“early invaders”) were decreased (Figure 5, F). Hence, we concluded that the pathway of host cell entry directly affects phagosomal escape.

Added in the revised version: We next monitored phagosomal escape over 10 h in cells pretreated with ASM inhibitors (Figure 6, A; Supp. Video 5; see Supp. Figure 7, for absolute bacteria numbers of individual replicates). ASM inhibition enhanced escape rates 3 h p.i. when compared to untreated controls. In contrast, 6 h p.i. escape rates were found decreased upon ASM inhibition.

Blocking ASM-dependent internalization affects intracellular replication and host cell survival.

Earlier version: (A, B) HeLa cells expressing RFP-CWT and LyseninW20A-YFP were treated with amitriptyline or ARC39 or were left untreated (Ctrl). The cells were subsequently infected with S. aureus and infection was monitored for 10 h by CLSM. Proportion of escaped bacteria (A) or the intracellular replication (B) was determined (n= 5). (C) HuLEC were stained with BODIPY-FL-C12-SM (green), and either were pretreated with 100 ng/ml of β-toxin or were left untreated. Then, cells were infected with S. aureus (red) and infection was monitored by CLSM for 19 h. (D, E) HuLEC were treated with amitriptyline or β-toxin and subsequently infected with S. aureus. After 21 h, plasma membrane integrity was measured by 7-AAD staining (D), and proportion of apoptotic cells was determined by annexin V staining (E) (n=5). See also Supp. Figure 6. (F) Model: S. aureus interacts with an unknown primary receptor (1) that triggers production of NAADP, which in turn activates TPC1 to mediate lysosomal Ca2+ release and activation of Syt7 (2). Syt7-dependent lysosomal exocytosis (3) leads to release of ASM (4) and production of Cer on the plasma membrane (5) thereby resulting in the recruitment of co-receptors (6) and rapid internalization of S. aureus by host cells (7). Bacteria that enter the host cells via the rapid pathway experience a different intracellular fate than bacteria that entered via other pathways. Statistics: Mixed effects analysis (REML) and Tukey‘s multiple comparison test. Graphs represent mean ± SD. *p≤0.05.

Revised version: (A, B) HeLa cells expressing RFP-CWT and LyseninW20A-YFP were treated with amitriptyline or ARC39 or were left untreated (Ctrl). The cells were subsequently infected with S. aureus and infection was monitored for 10 h by CLSM. Proportion of escaped bacteria (A) or the intracellular replication (B) was determined (n= 5). (C) HuLEC were stained with BODIPY-FL-C12-SM (green), and were either pretreated with 100 ng/ml of β-toxin or were left untreated. Then, cells were infected with S. aureus (red) and infection was monitored by CLSM for 19 h. (D, E) HuLEC were treated with amitriptyline or β-toxin and subsequently infected with S. aureus. After 21 h, plasma membrane integrity was measured by 7-AAD staining (D), and proportion of apoptotic cells was determined by annexin V staining (E) (n=5). See also Supp. Figure 8. (F) Model: S. aureus interacts with an unknown primary receptor (1) that triggers production of NAADP, which in turn activates TPC1 to mediate lysosomal Ca2+ release and activation of Syt7 (2). Syt7-dependent lysosomal exocytosis (3) leads to release of ASM (4) and production of Cer on the plasma membrane (5) thereby resulting in the recruitment of co-receptors [6,(83)] and rapid internalization of S. aureus by host cells (7). Bacteria, which enter the host cells via the rapid pathway, experience a different intracellular fate than bacteria that employed other pathways. Statistics: Mixed effects analysis (REML) and Tukey‘s multiple comparison test. Graphs represent mean ± SD. *p≤0.05. Created in BioRender.

Removed in the revised version: Created with BioRender.com/t33f620.

Removed in the revised version: Phagosomal escape of S. aureus thus is independent of SM levels within the infected phagosomes, but rather reflects the presence of SM and ASM at the plasma membrane during host cell entry.

Removed in the revised version: To demonstrate that phagosomal escape is dependent on the pathway of cell entry, we infected reporter cells with S. aureus for either a 10 min (“early invaders”) or a 30 min infection pulse (“early and late invaders”). Subsequently, we determined phagosomal escape rates (Figure 4, H) as well as the number of invaded bacteria (Supp. Figure 4, K). Independent of the multiplicity of infection (MOI) used, bacteria taken up within the first 10 min demonstrated significantly lower escape rates when compared to bacteria that were cocultured with host cells for 30 min.

Removed in the revised version: This was corroborated by HeLa YFP-CWT infected with S. aureus expressing a fluorescent protein (e.g. mRFP) for 30 min (“early and late invaders”). After 20 min, we initiated a second infection pulse for 10 min (“early invaders”) with a S. aureus strain expressing another fluorescence protein (e.g. Cerulean) and determined phagosomal escape rates of both bacterial recombinants (Figure 4, I). Again, early invaders demonstrated lower escape rates when compared to the 30 min control. This was independent of the fluorescence expressed by the strain that was used for each infection pulse.

Removed in the revised version: We next monitored phagosomal escape over 10 h in cells pretreated with ASM inhibitors (Figure 5, A). The compounds enhanced escape rates 3 h p.i. when compared to untreated controls. In contrast, 6 h p.i. escape rates were found decreased upon ASM inhibition.

Taken together, the efficiency of phagosomal escape of S. aureus is directly affected by the mode of cell entry, whereby the rapid SM/ASM-dependent pathway results in delayed phagosomal escape, while bacteria that entered host cells ASM-independently escape earlier during infection.

Earlier version: Since translocation to the host cytosol is a prerequisite for S. aureus replication in non-professional phagocytes, we recorded intracellular growth of S. aureus. In untreated controls replication started about 5 h p.i., and thus shortly after phagosomal escape (3-4 h p.i., Figure 5, B). By contrast, we did not observe significant bacterial replication in cells treated with ASM inhibitors (Figure 5, B).

Revised version: Since translocation to the host cytosol is a prerequisite for S. aureus replication in non-professional phagocytes, we recorded intracellular growth of S. aureus. In untreated controls replication started about 5 h p.i., and thus shortly after phagosomal escape (3-4 h p.i., Figure 6, B). By contrast, we did not observe significant bacterial replication in cells treated with ASM inhibitors (Figure 6, B).

Earlier version: We next treated HuLEC with β-toxin, stained the cells with BODIPY-FL-C12-SM and infected the cells with fluorescent S. aureus. Live cell imaging demonstrated that most untreated host cells died after 11 h p.i., whereas β-toxin-treated samples showed higher survival rates (Figure 5, C). This was also observed by cell death assays measuring i) membrane integrity (7-AAD staining, Figure 5, D), ii) apoptosis (Annexin V staining, Figure 5, E), iii) number of host cells that remained attached to the substratum (Supp. Figure 6, A) and iv) host cell lysis by LDH release (Supp. Figure 6, B).

Revised version: We next treated HuLEC with β-toxin, stained the cells with BODIPY-FL-C12-SM and infected the cells with fluorescent S. aureus. Live cell imaging demonstrated that most untreated host cells died after 11 h p.i., whereas β-toxin-treated samples showed higher survival rates (Figure 6, C; Supp. Video 6). This was also observed by cell death assays measuring i) membrane integrity (7-AAD staining, Figure 6, D), ii) apoptosis (Annexin V staining, Figure 6, E), iii) number of host cells that remained attached to the substratum (Supp. Figure 8, A) and iv) host cell lysis by LDH release (Supp. Figure 8, B) at 21 h p.i..

Discussion

Earlier version: ASM previously has been implicated in the invasion of several viral, bacterial, and eukaryotic pathogens (5–11). However, the dynamics of ASM-mediated host cell entry, the underlying receptors as well as effects on post-invasion events are only poorly understood.

Revised version: ASM has been previously reported to be implicated in the invasion of several viral, bacterial, and eukaryotic pathogens (7, 8, 13-17). However, the dynamics of ASM-mediated host cell entry, the underlying receptors as well as effects on post-invasion events are only poorly understood.

Earlier version: We here show that ASM is involved within minutes after the pathogen contacts host cells and is concurrent with ASM-independent uptake pathways. We demonstrate that the infection outcome of bacteria that enter the host cells via this rapid pathway is distinct from other entry mechanisms.

Revised version: We here show that ASM is involved in the uptake of S. aureus by certain human cells within minutes after the pathogen contacts the host cell and this pathway is concurrent with ASM-independent uptake pathways. We demonstrate that the infection outcome of bacteria that enter the host cells via this rapid pathway is distinct from other entry mechanisms.

Added in the revised version: We observed a reduction of bacterial invasiveness upon treatment with different ASM inhibitors (Figure 2, G) or removal of the enzyme substrate SM from the plasma membrane (Figure 2, L) confirming that ASM activity on the host cell surface is crucial for pathogen internalization. Surprisingly, the uptake of S. aureus by ASM K.O. cells was similar to wildtype cells (Figure 2, I). However, in ASM K.O.s S. aureus invasion was only slightly affected by the ASM inhibitor ARC39, while ARC39-treated wildtype cells showed a stronger reduction (Figure 2, J). We concluded that reduced invasion in wildtype cells upon ARC39 treatment is mostly due to inhibition of ASM and consistently, the effect of ARC39 is strongly attenuated in ASM K.O.s. Hence, we postulate that other ASM-independent host cell pathways are upregulated in ASM K.O.s, which might be caused by the strongly altered sphingolipid profile detected in ASM K.O.s (Figure 2, K). For instance, S. aureus internalization by host cells is strongly restricted by caveolin-1, a protein associated with lipid membrane domains (76, 77). Accordingly, alterations in membrane composition might affect distribution of caveolin-1 on the plasma membrane, thereby fostering uptake of S. aureus. In consistence, previous studies reported dysfunction of caveolin-1-dependent endocytic processes in ASM-deficient cells (78, 79).

Added in the revised version: Release of ASM is caused by lysosomal exocytosis (6). Conventional methods for detection of lysosomal exocytosis often depend on the usage of anti-LAMP1 antibodies that are added to the cell culture medium for monitoring LAMP1 exposure on the plasma membrane (66). However, most antibodies are non-specifically bound by staphylococcal protein A (84) and hence, these protocols are not compatible with S. aureus infection studies. Therefore, we developed a novel approach to measure lysosomal exocytosis via a split luciferase-based assay that demonstrated lysosomal exocytosis during S. aureus infection (Figure 2, A-C). Consequently, depleting host cells from lysosomes that can undergo Ca2+-dependent exocytosis by ionomycin pretreatment, reduced invasion, even though ionomycin can influence general host cells Ca2+ homeostasis (85). Additionally, the lysosomal exocytosis inhibitor Vacuolin-1 resulted in drastically reduced internalization of bacteria (Figure 2, D), although the substance is known for side effects such as inhibition of PIKfyve thereby affecting the endo-lysosomal compartment (86). However, we confirmed the role of lysosomal exocytosis for S. aureus invasion of host cells by genetic ablation of Syt7, a key protein in lysosomal exocytosis (69).) It remains elusive, which other factors are involved in delivery of lysosomes to and fusion of lysosomes with the plasma membrane. Next to Syt7, there might be other proteins such as Synaptotagmin 1 with overlapping function that could compensate for the absence of Syt7 (87).

Removed in the revised version: We observed a reduction of bacterial invasiveness upon treatment with different ASM inhibitors (Figure 2, C, D) or removal of the enzyme substrate SM from the plasma membrane (Figure 2, E) confirming that ASM activity on the host cell surface is crucial for pathogen internalization.

Removed in the revised version: Release of ASM is caused by lysosomal exocytosis (56). Consequently, the inhibitor Vacuolin-1 as well as genetic ablation of Syt7, a key protein in lysosomal exocytosis (55), resulted in drastically reduced internalization of bacteria (Figure 2, A, B).

Earlier version: Elevation of cytosolic Ca2+ levels lead to recruitment of lysosomes to the cell surface (53). Whereas most studies implicated Ca2+ influx from the extracellular milieu (8, 10), we here demonstrate that S. aureus mainly triggers lysosomal Ca2+ mobilization as was shown by treatment with the inhibitor trans-Ned19 as well as a genetic K.O. of the endo-lysosomal Ca2+ channel TPC1 (Figure 1, D and E; Supp. Figure 1, G). TPCs have previously been associated with regulation of phagocytic processes (62) and host cell entry of several viruses such as Ebola virus (63), Middle East respiratory syndrome coronavirus [MERS-CoV, (64)] or SARS-CoV-II (65).

Revised version: Elevation of cytosolic Ca2+ levels lead to recruitment of lysosomes to the cell surface (67). Whereas most studies implicated Ca2+ influx from the extracellular milieu (7, 16), we here demonstrate that S. aureus mainly triggers lysosomal Ca2+ mobilization as was shown by treatment with the inhibitor trans-Ned19 as well as a genetic K.O. of the endo-lysosomal Ca2+ channel TPC1 (Figure 1, B and E; Supp. Figure 1, C). There are two TPCs (TPC1 and TPC2) that mediate lysosomal Ca2+ liberation (88), which have previously been associated with regulation of phagocytic processes (89) and host cell entry of several viruses such as Ebola virus (90), Middle East respiratory syndrome coronavirus [MERS-CoV, (91)] or SARS-CoV-II (92). In line with our study, absence of one of the two TPCs severely interfered with infection (90)(91).

Earlier version: TPC-dependent Ca2+ liberation was shown to induce exocytosis, for instance, during release of cytolytic granules from T cells (66) or in interplay with TRPM4 during secretion of insulin (67). Consistently, the TRPM4 channel blocker 9-Phenantrol reduced internalization of S. aureus by host cells (Figure 1, B). Interestingly, exogenous addition of a bacterial SMase immediately before infection rescued the invasion defect in TPC1 and Syt7 K.O. cells (Figure 2, F) indicating that reduced invasion results from a decreased delivery of ASM to the host cell surface. Thus, NAADP-dependent Ca2+ liberation from lysosomal stores mediates Syt7-dependent fusion of lysosomes with the plasma membrane resulting in the release of ASM (Figure 5, F).

Revised version: TPC-dependent Ca2+ liberation was shown to induce exocytosis, for instance, during release of cytolytic granules from T cells (93) or during secretion of insulin (94). Our observations suggest that TPC1 is also involved in exocytosis of lysosomes, since exogenous addition of a bacterial SMase immediately before infection rescued the invasion defect in TPC1 and Syt7 K.O. cells (Figure 2, M) indicating that reduced invasion results from a decreased delivery of ASM to the host cell surface. Thus, NAADP-dependent Ca2+ liberation from lysosomal stores mediates Syt7-dependent fusion of lysosomes with the plasma membrane resulting in the release of ASM (Figure 6, F). While ionomycin treatment led to exocytosis of all “releasable” lysosomes, infection with S. aureus led to lysosomal exocytosis of only ∼30% of the ionomycin control (Figure 2, C), suggesting that S. aureus likely triggers local lysosomal exocytosis predominantly at host-bacteria contact sites, whereas ionomycin acts globally on all cells. A similar scenario has been observed for exocytosis of cytolytic granules in T lymphocytes, which differed between TPC/NAADP- and ionomycin-induction (93).

Added in the revised version: The conversion of SM to ceramide by ASM has been associated with the generation of ceramide-enriched platforms. These represent lipid microdomains which facilitate the recruitment of certain protein receptors (95, 96). We previously observed the interaction of S. aureus with several host cell surface receptors such as NRCAM or PTK7. The interaction of S. aureus with several host receptors was reduced upon blocking ASM-dependent internalization of S. aureus (83). It is thus tempting to speculate that accessory receptors identified in this study may be recruited to bacteria-host contact sites to facilitate rapid pathogen internalization.

Earlier version: Inhibition of CD38, a known NAADP producer (49, 50), did not influence uptake of staphylococci (Supp. Figure 1, L). By contrast, we detected a decreased invasion in cells deficient for the Toll-like receptor adaptor protein SARM1, an alternative NAADP producer (52). However, the role of SARM1 in lysosomal exocytosis and ASM release requires further investigation. The impact of SARM1 on internalization of S. aureus was rather moderate, suggesting the existence of other NAADP-generating enzymes.

Revised version: Inhibition of CD38, a known NAADP producer (56, 57), did not influence uptake of staphylococci (Supp. Figure 1, H). By contrast, we detected a decreased invasion in cells deficient for the Toll-like receptor adaptor protein SARM1, an alternative NAADP producer [Figure 1, D; (59)]. However, the role of SARM1 in lysosomal exocytosis and ASM release requires further investigation. The impact of SARM1 on internalization of S. aureus was rather moderate, suggesting the existence of other NAADP-generating enzymes, such as the recently identified NADPH oxidases DUOX1, DUOX2, NOX1 and NOX2(60).

Earlier version: One important characteristic of the ASM-dependent internalization pathway is its velocity - taking place within the initial ten minutes of an infection with S. aureus, during which treatment with small molecules (Figure 3, D), as well as genetic ablation of TPC1 (Figure 1, E) and Syt7 (Figure 2, B) resulted in reduction of intracellular bacteria. This was accompanied by the quick formation of S. aureus-containing phagosomes decorated with fluorescent SM analogs (Figure 3, A; Supp. Figure 3, Supp. Video 1, Supp. Video 2).

Revised version: One important characteristic of the ASM-dependent internalization pathway is its velocity - taking place within the initial ten minutes of an infection with S. aureus, during which treatment with small molecules (Figure 3, D), as well as genetic ablation of TPC1 (Figure 1, C) and Syt7 (Figure 2, F) resulted in reduction of intracellular bacteria. This was accompanied by the quick formation of S. aureus-containing phagosomes decorated with fluorescent SM analogs (Figure 3, A; Supp. Figure 3, Supp. Video 1, Supp. Video 2).

Added in the revised version: Neither treatment with inhibitors or enzymes nor genetic ablation blocking the rapid internalization pathway abolished the bacterial invasion of host cells completely. This supports our hypothesis that several concomitant host cell entry pathways exist, which is corroborated by a large number of host cell receptors interacting with S. aureus (13, 28, 31, 33-38). This pathway exists only in certain host cell types, since among the cell types, as is -for instance-illustrated by our experiments in which internalization of S. aureus was only sensitive to ASM inhibition in endothelial (HuLEC and HuVEC) and HeLa cells (Figure 2, G). The proportion of bacteria taken up by rapid internalization is dependent on the duration of infection and is higher in early (∼50-80% of all internalized bacteria 10 min p.i.) compared to late infection (∼15-50% of all internalized bacteria 30 min p.i.).

Added in the revised version: For simplicity we describe the pathway investigated in this study as “rapid” or “ASM-dependent”, however a clear cut between this and other pathways cannot be easily delineated. Additionally, internalization factors such as host cell receptors might possess overlapping function and could be involved in several pathways. In this context, ASM might simply accelerate certain internalization pathways, which could also take place in a slower fashion without involvement of ASM.

Earlier version: When we blocked the rapid internalization pathway by ASM inhibitors (Figure 4, A, C; Supp. Figure 4, B), Vacuolin-1 (Supp. Figure 4, F) or by removal of surface SM (Figure 4, A, E; Supp. Figure 4, B, H), we observed changes in both, maturation of bacteria-containing vesicles and phagosomal escape. A delayed maturation of phagosomes may affect acidification of the vesicles, a property that is sensed by S. aureus leading to expression of a different subset of virulence factors (68) and hence may alter intracellular pathogenicity. Phagosomal escape of S. aureus only was dependent on SM on the plasma membrane during invasion but not on the presence of SM within the phago-endosome (Figure 4, D-G; Supp. Figure 4, H, I). This not only corroborated that phagosomal escape of clinically relevant S. aureus strains is independent of the bacterial SMase β-toxin (36), but also demonstrated that the outcome of phagosomal escape is influenced by processes at the plasma membrane during cell entry.

Revised version: When we blocked the rapid internalization pathway by ASM inhibitors (Figure 4, A, C; Supp. Figure 4, B), Vacuolin-1 (Figure 4, E, Supp. Figure 4) or by removal of surface SM (Figure 5, B; Supp. Figure 4, B, C), we observed changes in both, maturation of bacteria-containing vesicles and phagosomal escape. A delayed maturation of phagosomes may affect acidification of the vesicles, a property that is sensed by S. aureus leading to expression of a different subset of virulence factors (97) and hence may alter intracellular pathogenicity.

Added in the revised version: The bacterial SMase β-toxin originates from S. aureus, however, most human pathogenic S. aureus strains including JE2, the strain we use in the present study, do not produce β-toxin, due to a genomic integration of a prophage (80). Phagosomal escape of S. aureus only was dependent on SM on the plasma membrane during invasion but not on the presence of SM within the phago-endosome (Figure 5, A-C; Supp. Figure 4, I, J). This not only corroborated that phagosomal escape of clinically relevant S. aureus strains is independent of the bacterial SMase β-toxin (42), but also demonstrated that the outcome of phagosomal escape is influenced by processes at the plasma membrane during cell entry.

Earlier version: To exclude that inhibitor or bSMase treatments caused changes in phagosomal escape by affecting cellular processes other than bacterial uptake, we infected untreated WT cells with S. aureus for 10 or 30 min, respectively, and monitored phagosomal escape. Bacteria that invaded within the first 10 minutes after contacting host cells (“early invaders”) are predominantly ASM-dependent (Figure 3) and demonstrated significantly less phagosomal escape (Figure 4, G, H). Consistently, blocking the ASM-dependent pathway led to increased proportions of ASM-independent invaders, which in turn resulted in higher escape rates (Figure 4, C, E; Supp. Figure 4, F). Intracellular S. aureus infections thus are directly influenced by the mode of internalization. This ultimately resulted in lower host cytotoxicity (Figure 5, C-E; Supp. Figure 6) and bacterial replication (Figure 5, B) when we blocked the ASM-dependent invasion.

Revised version: To exclude that inhibitor or bacterial SMase treatment caused changes in phagosomal escape by affecting cellular processes other than bacterial uptake, we infected untreated wildtype cells with S. aureus for 10 or 30 min, respectively, and monitored phagosomal escape. Bacteria that invaded within the first 10 minutes after contacting host cells (“early invaders”) are predominantly ASM-dependent (Figure 3) and demonstrated significantly less phagosomal escape (Figure 5, G, H). Consistently, blocking the ASM-dependent pathway led to increased proportions of ASM-independent invaders, which in turn resulted in higher escape rates (Figure 4, C-E; Supp. Figure 4, G). Intracellular S. aureus infections thus are directly influenced by the mode of internalization. This ultimately resulted in lower host cytotoxicity (Figure 6, C-E; Supp. Figure 8) and bacterial replication (Figure 6, B) when we blocked the ASM-dependent invasion.

Earlier version: An uptake-related intracellular fate has been suggested for Mycobacterium bovis (39), Toxoplasma gondii (41) or Brucella abortus (42). However, in these studies the host cells were either depleted from cholesterol, which causes a variety of side effects (40), or internalization was artificially induced by chemicals or ectopic expression of receptors, respectively.

Revised version: An uptake-related intracellular fate has been suggested for Mycobacterium bovis (46), Toxoplasma gondii (48) or Brucella abortus (49). However, in these studies the host cells were either depleted from cholesterol, which causes a variety of side effects (47), or internalization was artificially induced by chemicals or ectopic expression of receptors, respectively.

Earlier version: Taken together, we here describe a rapid internalization pathway for S. aureus into human epithelial and endothelial cells. Bacteria that enter the host cells via this infection route face a distinct intracellular fate. Here we show, to our knowledge, for the first time that host cells within the same population can be invaded by a pathogen via distinct routes and that this results in a differential infection outcome. Rapid internalization may prove beneficial for pathogens in an in vivo setting since internalization by host cells shortens the exposure to the innate immune system within a host. Interestingly, the absence of ASM in Smpd1-/- mice enhanced the potency of antibiotics to clear S. aureus sepsis (69). It is tempting to speculate that reduced host cell invasion in absence of ASM caused a prolonged exposure of the pathogen to antibiotics. Several FIASMAs are already approved drugs (70), and may have future use in infection treatment.

Revised version: Taken together, we here describe a rapid internalization pathway for S. aureus into human epithelial and endothelial cells. Bacteria that enter the host cells via this infection route face a distinct intracellular fate. Here we show, to our knowledge, for the first time that host cells within the same population can be invaded by a pathogen via distinct routes and that this results in a differential infection outcome. Rapid internalization may prove beneficial for pathogens in an in vivo setting since internalization by host cells shortens the exposure to the innate immune system within a host. Interestingly, the absence of ASM in Smpd1-/- mice enhanced the potency of antibiotics to clear S. aureus sepsis (98). It is tempting to speculate that reduced host cell invasion in absence of ASM caused a prolonged exposure of the pathogen to antibiotics. Several FIASMAs are already approved drugs (99), and may have future use in infection treatment.